Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# SECTION 2 ONLY
Name = "Thanakrit Samoena"

Student ID = "67070060"

# <mark>Mock Final Exam เพื่อเตรียมความพร้อมก่อนสอบ (ได้คะแนนเต็มทุกคน (ถ้าส่ง)) </mark>

- <mark>**Write you codes after `# YOUR CODE HERE`**</mark>
- <mark>**ห้าม Import library อื่น ๆ นอกจาก `numpy`**</mark>

```python
import numpy as np
```

---
## (5 คะแนน) <mark>Exercise #1 (Easy): **นับเลข คู่ (Even) และ คี่ (Odd)** จาก `list` ตัวเลขที่กำหนด</mark>

- เขียนฟังก์ชัน `count_even_odd(numbers)` โดยมีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับพารามิเตอร์เป็น `list` ของตัวเลข (`int`)
  2) นับจำนวนเลขคู่ (even) และเลขคี่ (odd) ใน `list`
  3) Return ผลลัพธ์เป็น `dict` ที่มี **key** `'even'` และ `'odd'` แสดง***จำนวน***เลขคู่และเลขคี่ตามลำดับ

- เช่น
    - `count_even_odd([1, 2, 3, 4, 5])` จะ return `{'even': 2, 'odd': 3}`
    - `count_even_odd([10, 15, 20, 25, 30])` จะ return `{'even': 3, 'odd': 2}`

<!-- grade_id: cell-b385a4f87b16abdc, points: 0, grade: false, solution: true -->
```python
def count_even_odd(numbers):
    # YOUR CODE HERE
    num = {'even' : 0, 'odd' : 0}
    for i in numbers:
        if i % 2 == 0:
            num['even'] += 1
        else:
            num['odd'] += 1
    return num
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-70b2242df256dab7, points: 0.5, grade: true, solution: false -->
```python
# Case 1
assert count_even_odd([1, 2, 3, 4, 5]) == {'even': 2, 'odd': 3}
```

<!-- grade_id: cell-62ef87c26b10a9df, points: 0.5, grade: true, solution: false -->
```python
# Case 2
assert count_even_odd([10, 15, 20, 25, 30]) == {'even': 3, 'odd': 2}
```

<!-- grade_id: cell-aedd81c194858443, points: 0.5, grade: true, solution: false -->
```python
# Case 3
assert count_even_odd([0, -1, -2, -3, -4]) == {'even': 3, 'odd': 2}
```

<!-- grade_id: cell-0399901ffb0f963a, points: 0.5, grade: true, solution: false -->
```python
# Case 4
assert count_even_odd([5, 7, 9]) == {'even': 0, 'odd': 3}
```

<!-- grade_id: cell-3e94dd5f44647502, points: 0.5, grade: true, solution: false -->
```python
# Case 5
assert count_even_odd([2, 4, 6, 8]) == {'even': 4, 'odd': 0}
```

<!-- grade_id: cell-e5aa41b3213e22ad, points: 0.5, grade: true, solution: false -->
```python
# Case 6
assert count_even_odd([1,2,-9,3,5,8,22,-6,4,3]) == {'even': 5, 'odd': 5}
```

<!-- grade_id: cell-b665e0bf2c00ba49, points: 0.5, grade: true, solution: false -->
```python
# Case 7
assert count_even_odd([-2, 1, 3, 4]) == {'even': 2, 'odd': 2}
```

<!-- grade_id: cell-0802a5ac44c9374f, points: 0.5, grade: true, solution: false -->
```python
# Case 8
assert count_even_odd([2, 5, 8, 13, 21, 34]) == {'even': 3, 'odd': 3}
```

<!-- grade_id: cell-0c8d5c2edec55249, points: 0.5, grade: true, solution: false -->
```python
# Case 9
assert count_even_odd([100, 200, 300]) == {'even': 3, 'odd': 0}
```

<!-- grade_id: cell-f7e8a9b8794e6674, points: 0.5, grade: true, solution: false -->
```python
# Case 10
assert count_even_odd([-10, -5, -20, 0]) == {'even': 3, 'odd': 1}
```

---
## (5 คะแนน) <mark>Exercise #2 (Easy): **วิเคราะห์ตัวเลขจาก Text ไฟล์ที่่กำหนด**</mark>

- เขียนฟังก์ชัน `analyze_integer_file(filename)` โดยมีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับพารามิเตอร์เป็นชื่อไฟล์ (string) ที่ต้องการอ่าน
  2) เปิดไฟล์และอ่านจำนวนเต็มทุกบรรทัดจากไฟล์
  3) ฟังก์ชันจะทำการ return ค่าต่าง ๆ ในรูปของ `list` ดังนี้
     - ค่า max
     - ค่า min
     - จำนวนตัวเลข

- เช่นหากไฟล์ `integers.txt` มีตัวเลข `10 20 30 40 50` ฟังก์ชัน `analyze_integer_file('integers.txt')` จะ return เป็น `[10, 50, 5]`

<!-- grade_id: cell-1d0572a17c6c1c83, points: 0, grade: false, solution: true -->
```python
def analyze_integer_file(filename):
    # YOUR CODE HERE
    li = []
    with open(filename , mode="r") as f:
        for i in f:
          li.append(int(i))
    print(li)
    return [min(li),max(li),len(li)]       
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-5f2011a95bf9daa0, points: 1, grade: true, solution: false -->
```python
# Case 1
assert analyze_integer_file('number_1.txt') == [5, 35, 11]
```

Output:
```
[22, 35, 18, 27, 14, 30, 10, 5, 9, 24, 21]
```

<!-- grade_id: cell-7646ee386a204c6f, points: 1, grade: true, solution: false -->
```python
# Case 2
assert analyze_integer_file('number_2.txt') == [1, 29, 13]
```

Output:
```
[4, 15, 23, 7, 1, 12, 29, 3, 19, 10, 22, 17, 16]
```

<!-- grade_id: cell-2aee25a8fd295ae1, points: 1, grade: true, solution: false -->
```python
# Case 3
assert analyze_integer_file('number_3.txt') == [5, 50, 11]
```

Output:
```
[45, 32, 19, 50, 11, 25, 38, 14, 5, 8, 27]
```

<!-- grade_id: cell-2078565bbcc614ca, points: 1, grade: true, solution: false -->
```python
# Case 4
assert analyze_integer_file('number_4.txt') == [1, 17, 12]
```

Output:
```
[6, 13, 9, 16, 2, 8, 17, 10, 12, 15, 4, 1]
```

<!-- grade_id: cell-debe0ab68eb58de0, points: 1, grade: true, solution: false -->
```python
# Case 5
assert analyze_integer_file('number_5.txt') == [1, 19, 10]
```

Output:
```
[5, 12, 19, 3, 7, 14, 8, 11, 1, 6]
```

---
## (5 คะแนน)  <mark>Exercise #3 (Easy): **แยกเลขคู่ (Even) หรือ คี่ (Odd) จาก `list` ที่กำหนด**</mark>

- เขียนฟังก์ชัน `extract_values()` ที่รับพารามิเตอร์ 2 ค่า คือ 
  1) `list` ของเลข integers
  2) สตริงที่ระบุเงื่อนไขการแยกกลุ่มตัวเลข `'even'` หรือ `'odd'`  
- ฟังก์ชันนี้จะทำการแยกตัวเลขทั้งหมดที่เป็น **เลขคู่ (even)** หรือ **เลขคี่ (odd)** ตามที่ผู้ใช้เลือก ตามเงื่อนไขดังนี้
    1) หากผู้ใช้เลือก `'even'` ให้ return `list` ที่ประกอบด้วยตัวเลขคู่ทั้งหมดจาก `list` ของ input
    2) หากผู้ใช้เลือก `'odd'` ให้ return `list` ที่ประกอบด้วยตัวเลขคี่ทั้งหมดจาก `list` ของ input
    3) หากพารามิเตอร์เป็นค่าว่าง หรือเลือกค่าอื่นที่ไม่ใช่ `'even'` หรือ `'odd'` ให้คืนสตริงว่า `"Invalid selection."`
    
- ตัวอย่างเช่น:
    - `extract_values([1, 2, 3, 5, 7, 8, 9], 'even')` จะคืนค่า `[2, 8]`
    - `extract_values([1, 2, 3, 5, 7, 8, 9], 'odd')` จะคืนค่า `[1, 3, 5, 7, 9]`
    - `extract_values([1, 2, 3, 5, 7, 8, 9], 'invalid')` จะคืนค่า `"Invalid selection."`

<!-- grade_id: cell-a60b33c3768b3c6e, points: 0, grade: false, solution: true -->
```python
def extract_values(original_list, selection):
    # YOUR CODE HERE
    ans = []
    if selection == "even":
        for i in original_list:
            if i % 2 == 0:
                ans.append(i)
        return ans
    elif selection == "odd":
        for i in original_list:
            if i % 2:
                ans.append(i)
        return ans
    else:
        return "Invalid selection."
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-aeb5df92b8fbddf6, points: 0.5, grade: true, solution: false -->
```python
# Case 1
assert extract_values([1, 2, 3, 5, 7, 8, 9], 'even') == [2, 8]
```

<!-- grade_id: cell-b49c116608a72628, points: 0.5, grade: true, solution: false -->
```python
# Case 2
assert extract_values([1, 2, 3, 5, 7, 8, 9], 'odd') == [1, 3, 5, 7, 9]
```

<!-- grade_id: cell-310f13f5d3bd6449, points: 0.5, grade: true, solution: false -->
```python
# Case 3
assert extract_values([1, 2, 3, 5, 7, 8, 9], 'invalid') == "Invalid selection."
```

<!-- grade_id: cell-a01dcd26e2e1e288, points: 0.5, grade: true, solution: false -->
```python
# Case 4
assert extract_values([2, 4, 6, 8, 10, 12, 14, 16, 18, 20], 'even') == [2, 4, 6, 8, 10, 12, 14, 16, 18, 20]
```

<!-- grade_id: cell-4dd17b2e440d5b17, points: 0.5, grade: true, solution: false -->
```python
# Case 5
assert extract_values([1, 1, 1, 1, 1, 1, 1, 1], 'even') == []
```

<!-- grade_id: cell-aa10302450069f7a, points: 0.5, grade: true, solution: false -->
```python
# Case 6:
assert extract_values([42], 'even') == [42]
```

<!-- grade_id: cell-8cc805557f9611e8, points: 0.5, grade: true, solution: false -->
```python
# Case 7
assert extract_values([-10, -3, -5, 6, 7, 0], 'odd') == [-3, -5, 7]
```

<!-- grade_id: cell-60c5d57e8d84cf1d, points: 0.5, grade: true, solution: false -->
```python
# Case 8
assert extract_values([-2, 2, -2, 2, -2, 2], 'odd') == []
```

<!-- grade_id: cell-2989513bf03a5b13, points: 0.5, grade: true, solution: false -->
```python
# Case 9
assert extract_values(list(range(20)), 'odd') == [1, 3, 5, 7, 9, 11, 13, 15, 17, 19]
```

<!-- grade_id: cell-d7fe54d03fd7d847, points: 0.5, grade: true, solution: false -->
```python
# Case 10
assert extract_values([100, 201, 300, 404, 505, 606, 707, 808, 909, 1000], 'even') == [100, 300, 404, 606, 808, 1000]
```

---
## (5 คะแนน)  <mark>Exercise #4 (Easy): **การจัดการ Error ด้วย Try/Except**</mark>

- เขียนฟังก์ชัน `divide_by_user_input()` ที่มีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับ input จำนวน 1 ตัว
  2) ทำการหาร input ด้วยจำนวน `100`
  3) หากผู้ใช้ป้อนค่า `0` ให้แสดงข้อความว่า `"Cannot divide by zero."`
  4) หากผู้ใช้ป้อนค่าที่ไม่ใช่จำนวนเต็ม (เช่น ตัวอักษร) ให้แสดงข้อความว่า `"Invalid input. Please enter a number."`
  
- เช่น
  - ถ้าผู้ใช้ป้อน `4` จะได้ผลลัพธ์ว่า `100 / 4 = 25.0`
  - ถ้าผู้ใช้ป้อน `0` จะได้ผลลัพธ์ว่า `"Cannot divide by zero."`
  - ถ้าผู้ใช้ป้อน `abc` จะได้ผลลัพธ์ว่า `"Invalid input. Please enter a number."`

<!-- grade_id: cell-b27bd436b8c16bf7, points: 0, grade: false, solution: true -->
```python
def divide_by_user_input(user_input):
    # YOUR CODE HERE
    try:  
        result = 100/user_input
    except TypeError:
        return "Invalid input. Please enter a number."
    except ZeroDivisionError:
        return "Cannot divide by zero."
    else:
        return f"100 / {user_input} = {result}"

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-a20d4ee6cbab7b2f, points: 0.5, grade: true, solution: false -->
```python
# Case 1
assert divide_by_user_input(4) == "100 / 4 = 25.0"
```

<!-- grade_id: cell-9cf58c902673425e, points: 0.5, grade: true, solution: false -->
```python
# Case 2
assert divide_by_user_input(10) == "100 / 10 = 10.0"
```

<!-- grade_id: cell-15aa7364f61f893d, points: 0.5, grade: true, solution: false -->
```python
# Case 3
assert divide_by_user_input(0) == "Cannot divide by zero."
```

<!-- grade_id: cell-3e6d40ea4de67ab3, points: 0.5, grade: true, solution: false -->
```python
# Case 4
assert divide_by_user_input('abc') == "Invalid input. Please enter a number."
```

<!-- grade_id: cell-2cb0c4faabb9fd1a, points: 0.5, grade: true, solution: false -->
```python
# Case 5
assert divide_by_user_input(-5) == "100 / -5 = -20.0"
```

<!-- grade_id: cell-58daac4286ff7059, points: 0.5, grade: true, solution: false -->
```python
# Case 6
assert divide_by_user_input(2.5) == "100 / 2.5 = 40.0"
```

<!-- grade_id: cell-e771579c771628a2, points: 0.5, grade: true, solution: false -->
```python
# Case 7
assert divide_by_user_input(1000) == "100 / 1000 = 0.1"
```

<!-- grade_id: cell-a7109663728e7d65, points: 0.5, grade: true, solution: false -->
```python
# Case 8
assert divide_by_user_input(1) == "100 / 1 = 100.0"
```

<!-- grade_id: cell-76aaed0c1329cdc2, points: 0.5, grade: true, solution: false -->
```python
# Case 9
assert divide_by_user_input(999) == "100 / 999 = 0.1001001001001001"
```

<!-- grade_id: cell-86f0113b4566067c, points: 0.5, grade: true, solution: false -->
```python
# Case 10
assert divide_by_user_input(10000000) == "100 / 10000000 = 1e-05"
```

---
## (5 คะแนน)  <mark>Exercise #5 (Easy): **การจัดการ String**</mark>

- เขียนฟังก์ชัน `reverse_and_upperCase()` ตามเงื่อนไขต่อไปนี้
  1) ฟังก์ชันนี้จะรับพารามิเตอร์เป็นสตริง 1 จำนวน
  2) ทำการกลับทิศทางของสตริง (reverse)
  3) แปลงสตริงที่กลับทิศทางแล้วให้เป็นตัวพิมพ์ใหญ่ (Upper Case)
  4) Return ค่าสตริงใหม่เป็นผลลัพธ์ของฟังก์ชัน
- เช่น
  - `reverse_and_upperCase('Hello World!')` จะได้ `'!DLROW OLLEH'`
  - `reverse_and_upperCase('Python')` จะได้ `'NOHTYP'`

<!-- grade_id: cell-7427492f5698053f, points: 0, grade: false, solution: true -->
```python
def reverse_and_upperCase(input_string):
    # YOUR CODE HERE
    return input_string[::-1].upper()
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-7822110936ae48b9, points: 0.5, grade: true, solution: false -->
```python
# Case 1
assert reverse_and_upperCase('Hello World!') == '!DLROW OLLEH'
```

<!-- grade_id: cell-b90a0a27a73743a5, points: 0.5, grade: true, solution: false -->
```python
# Case 2
assert reverse_and_upperCase('Python') == 'NOHTYP'
```

<!-- grade_id: cell-471c63a9045726ac, points: 0.5, grade: true, solution: false -->
```python
# Case 3
assert reverse_and_upperCase('') == ''
```

<!-- grade_id: cell-07c05c54b96220c9, points: 0.5, grade: true, solution: false -->
```python
# Case 4
assert reverse_and_upperCase('12345One-Two-three') == 'EERHT-OWT-ENO54321'
```

<!-- grade_id: cell-255f428eff183cf9, points: 0.5, grade: true, solution: false -->
```python
# Case 5
assert reverse_and_upperCase('abc!@#') == '#@!CBA'
```

<!-- grade_id: cell-e9d9b513bf8e361d, points: 0.5, grade: true, solution: false -->
```python
# Case 6
assert reverse_and_upperCase('PyThOn') == 'NOHTYP'
```

<!-- grade_id: cell-bbf3d7b0513e4f80, points: 0.5, grade: true, solution: false -->
```python
# Case 7
assert reverse_and_upperCase('This is a long string for testing.') == '.GNITSET ROF GNIRTS GNOL A SI SIHT'
```

<!-- grade_id: cell-2dcb9285d86953b3, points: 0.5, grade: true, solution: false -->
```python
# Case 8
assert reverse_and_upperCase('Saippuakauppias') == 'SAIPPUAKAUPPIAS'
```

<!-- grade_id: cell-51aac9cf588e2598, points: 0.5, grade: true, solution: false -->
```python
# Case 9
assert reverse_and_upperCase('   spaces   ') == '   SECAPS   '
```

<!-- grade_id: cell-3b0d3e8f0f6e362f, points: 0.5, grade: true, solution: false -->
```python
# Case 10
assert reverse_and_upperCase('Rotavator') == 'ROTAVATOR'
```

---
## (5 คะแนน)  <mark>Exercise #6 (Easy): **การใช้ `list` และ `tuple`**</mark>

- เขียนฟังก์ชัน `generate_and_convert(m, n)` ที่มีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับพารามิเตอร์สองตัว คือ `m` และ `n` ซึ่งเป็นจำนวนเต็มเพื่อสร้างตัวเลขตั้งแต่ `m` ถึง `n` ที่หารลงตัวด้วย `3` และ `5` ลงตัว
  2) return `list` และ `tuple`
- เช่น `generate_and_convert(1, 100)` จะ return `[15, 30, 45, 60, 75, 90]` และ `(15, 30, 45, 60, 75, 90)`

<!-- grade_id: cell-f52ea0527eb7134f, points: 0, grade: false, solution: true -->
```python
def generate_and_convert(m, n):
    # YOUR CODE HERE
    li3 = []
    for i in range(m,n+1):
        if not i%3 and not i%5:
            li3.append(i)
    print(li3,tuple(li3))
    return li3,tuple(li3)
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-1f6872fb6be290c8, points: 0.5, grade: true, solution: false -->
```python
# Case 1
assert generate_and_convert(1, 100) == ([15, 30, 45, 60, 75, 90], (15, 30, 45, 60, 75, 90))
```

Output:
```
[15, 30, 45, 60, 75, 90] (15, 30, 45, 60, 75, 90)
```

<!-- grade_id: cell-216b6100e973394d, points: 0.5, grade: true, solution: false -->
```python
# Case 2
assert generate_and_convert(1, 50) == ([15, 30, 45], (15, 30, 45))
```

Output:
```
[15, 30, 45] (15, 30, 45)
```

<!-- grade_id: cell-af1412f2dac40b8e, points: 0.5, grade: true, solution: false -->
```python
# Case 3
assert generate_and_convert(50, 150) == ([60, 75, 90, 105, 120, 135, 150], (60, 75, 90, 105, 120, 135, 150))
```

Output:
```
[60, 75, 90, 105, 120, 135, 150] (60, 75, 90, 105, 120, 135, 150)
```

<!-- grade_id: cell-417bd6be1dd61f51, points: 0.5, grade: true, solution: false -->
```python
# Case 4
assert generate_and_convert(1, 15) == ([15], (15,))  # Test with range 1 to 15
```

Output:
```
[15] (15,)
```

<!-- grade_id: cell-74f8c01edb46229d, points: 0.5, grade: true, solution: false -->
```python
# Case 5
assert generate_and_convert(30, 100) == ([30, 45, 60, 75, 90], (30, 45, 60, 75, 90))  # Test with range 30 to 100
```

Output:
```
[30, 45, 60, 75, 90] (30, 45, 60, 75, 90)
```

<!-- grade_id: cell-ecf81863e480bd41, points: 0.5, grade: true, solution: false -->
```python
# Case 6
assert generate_and_convert(1, 1000)[0] == [15, 30, 45, 60, 75, 90, 105, 120, 135, 150, 165, 180, 195, 210, 225, 240, 255, 270, 285, 300, 315, 330, 345, 360, 375, 390, 405, 420, 435, 450, 465, 480, 495, 510, 525, 540, 555, 570, 585, 600, 615, 630, 645, 660, 675, 690, 705, 720, 735, 750, 765, 780, 795, 810, 825, 840, 855, 870, 885, 900, 915, 930, 945, 960, 975, 990]
```

Output:
```
[15, 30, 45, 60, 75, 90, 105, 120, 135, 150, 165, 180, 195, 210, 225, 240, 255, 270, 285, 300, 315, 330, 345, 360, 375, 390, 405, 420, 435, 450, 465, 480, 495, 510, 525, 540, 555, 570, 585, 600, 615, 630, 645, 660, 675, 690, 705, 720, 735, 750, 765, 780, 795, 810, 825, 840, 855, 870, 885, 900, 915, 930, 945, 960, 975, 990] (15, 30, 45, 60, 75, 90, 105, 120, 135, 150, 165, 180, 195, 210, 225, 240, 255, 270, 285, 300, 315, 330, 345, 360, 375, 390, 405, 420, 435, 450, 465, 480, 495, 510, 525, 540, 555, 570, 585, 600, 615, 630, 645, 660, 675, 690, 705, 720, 735, 750, 765, 780, 795, 810, 825, 840, 855, 870, 885, 900, 915, 930, 945, 960, 975, 990)
```

<!-- grade_id: cell-cc1ebfabd21ee4d5, points: 0.5, grade: true, solution: false -->
```python
# Case 7
assert generate_and_convert(100, 200) == ([105, 120, 135, 150, 165, 180, 195], (105, 120, 135, 150, 165, 180, 195))  # Test with range 100 to 200
```

Output:
```
[105, 120, 135, 150, 165, 180, 195] (105, 120, 135, 150, 165, 180, 195)
```

<!-- grade_id: cell-776e2840e41317b8, points: 0.5, grade: true, solution: false -->
```python
# Case 8
assert generate_and_convert(-20, 10) == ([-15, 0], (-15, 0))
```

Output:
```
[-15, 0] (-15, 0)
```

<!-- grade_id: cell-c63b26604d5133dd, points: 0.5, grade: true, solution: false -->
```python
# Case 9
assert generate_and_convert(0, 30) == ([0, 15, 30], (0, 15, 30))
```

Output:
```
[0, 15, 30] (0, 15, 30)
```

<!-- grade_id: cell-921b10ad869d4b7d, points: 0.5, grade: true, solution: false -->
```python
# Case 10
assert generate_and_convert(10, 10) == ([], ())  # Test with no valid numbers in range
```

Output:
```
[] ()
```

---
## (15 คะแนน) <mark>Exercise #7 (Moderate): **จัดการข้อมูลหนังสือด้วย Dictionary**</mark>
- รายการหนังสือ `book_collection` เป็น `dict` ที่โครงสร้างดังนี้ `{'title': {'author': 'author name', 'year': year_number}}`
- เขียนฟังก์ชัน `manage_book_records(book_collection, action='', title='', author='', year=0)` ที่มีเงื่อนไขดังนี้
  1) ฟังก์ชันรับพารามิเตอร์เป็น `book_collection` `action` `title` `author` และ `year`
  2) หาก `action` เป็น `'add'` ให้ทำดังนี้
      - หากชื่อหนังสือที่ต้องการ `add` นั้นมีอยู่แล้ว ให้ทำการ return ข้อความว่า `"Book '{title}' is already in the collection."`
      - หากชื่อหนังสือที่ต้องการ `add` ยังไม่มีใน `book_collection` ทำการเพิ่มหนังสือลงใน `book_collection` และ return `book_collection` ออกจากฟังก์ชัน
  3) หาก `action` เป็น `'remove'` ให้ทำดังนี้
      - หากไม่มีชื่อหนังสือใน `book_collection` ให้ return ข้อความว่า `"Book '{title}' not found."`
      - หากมีชื่อหนังสือใน `book_collection` ให้ทำการลบหนังสือออกจาก `book_collection` และ return ข้อความว่า `"Book '{title}' has been removed."`
  4) หาก `action` เป็น `'modify'` ให้ทำดังนี้
      - หากไม่มีชื่อหนังสือใน `book_collection` ให้ return ข้อความว่า `"Book '{title}' not found."`
      - หากมีชื่อหนังสือใน `book_collection` ให้แก้ไขข้อมูลหนังสือที่มีอยู่ เช่น เปลี่ยน `title` หรือ `year` ตามที่กำหนด
- ตัวอย่างเช่น:
  - `manage_book_records(book_collection, 'add', 'The Great Gatsby', 'F. Scott Fitzgerald', 1925)` จะเพิ่มหนังสือใหม่
  - `manage_book_records(book_collection, 'remove', 'The Great Gatsby')` จะลบหนังสือและ return `"Book 'The Great Gatsby' has been removed."`
  - `manage_book_records(book_collection, 'modify', 'The Great Gatsby', year=2023)` จะปรับปรุงปีของหนังสือจาก `1925` เป็น `2003`

<!-- grade_id: cell-a5084a0054c5c8a2, points: 0, grade: false, solution: true -->
```python
book_collection = {}    # Initialize the book collection (empty dictionary)

def manage_book_records(book_collection, action='', title='', author='', year=0):
    # YOUR CODE HERE
    print(book_collection.keys())
    print(author)
    print(year)
    if action=="add":
        if title in book_collection.keys():
            return f"Book '{title}' is already in the collection."
        elif  title not in book_collection.keys():
            book_collection[title] = {"author" : author,"year" : year}
            return book_collection
    elif action == "remove":
        if title not in book_collection.keys():
            return f"Book '{title}' not found."
        else:
            del book_collection[title]
            return f"Book '{title}' has been removed."
    elif action == "modify":
        if title not in book_collection.keys():
            return f"Book '{title}' not found."
        else:
            if author:
                book_collection[title]['author'] = author
            if year:
                book_collection[title]['year'] = year
            return book_collection
            
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-ff25c184a9de6956, points: 1.5, grade: true, solution: false -->
```python
# Case 1
assert manage_book_records(book_collection, 'add', 'The Great Gatsby', 'F. Scott Fitzgerald', 1925) == {
    'The Great Gatsby': {'author': 'F. Scott Fitzgerald', 'year': 1925}
}
```

Output:
```
dict_keys([])
F. Scott Fitzgerald
1925
```

<!-- grade_id: cell-e9ef0b8af01ba836, points: 1.5, grade: true, solution: false -->
```python
# Case 2
assert manage_book_records(book_collection, 'add', '1984', 'George Orwell', 1949) == {
    'The Great Gatsby': {'author': 'F. Scott Fitzgerald', 'year': 1925},
    '1984': {'author': 'George Orwell', 'year': 1949}
}
```

Output:
```
dict_keys(['The Great Gatsby'])
George Orwell
1949
```

<!-- grade_id: cell-9e8b9b472c5b448d, points: 1.5, grade: true, solution: false -->
```python
# Case 3
assert manage_book_records(book_collection, 'add', 'The Great Gatsby', 'F. Scott Fitzgerald', 1925) == "Book 'The Great Gatsby' is already in the collection."
```

Output:
```
dict_keys(['The Great Gatsby', '1984'])
F. Scott Fitzgerald
1925
```

<!-- grade_id: cell-17633426c773ebc7, points: 1.5, grade: true, solution: false -->
```python
# Case 4
assert manage_book_records(book_collection, 'remove', '1984') == "Book '1984' has been removed."
```

Output:
```
dict_keys(['The Great Gatsby', '1984'])

0
```

<!-- grade_id: cell-f52d600d33abf719, points: 1.5, grade: true, solution: false -->
```python
# Case 5
assert manage_book_records(book_collection, 'remove', '1984') == "Book '1984' not found."
```

Output:
```
dict_keys(['The Great Gatsby'])

0
```

<!-- grade_id: cell-4754392771459f70, points: 1.5, grade: true, solution: false -->
```python
# Case 6
assert manage_book_records(book_collection, 'modify', 'The Great Gatsby', year=2023) == {
    'The Great Gatsby': {'author': 'F. Scott Fitzgerald', 'year': 2023}
}
```

Output:
```
dict_keys(['The Great Gatsby'])

2023
```

<!-- grade_id: cell-082edea77843d27e, points: 1.5, grade: true, solution: false -->
```python
# Case 7
assert manage_book_records(book_collection, 'modify', '1984', year=2023) == "Book '1984' not found."
```

Output:
```
dict_keys(['The Great Gatsby'])

2023
```

<!-- grade_id: cell-3152f5e40ac1d4b2, points: 1.5, grade: true, solution: false -->
```python
# Case 8
assert manage_book_records(book_collection, 'add', 'To Kill a Mockingbird', 'Harper Lee', 1960) == {
    'The Great Gatsby': {'author': 'F. Scott Fitzgerald', 'year': 2023},
    'To Kill a Mockingbird': {'author': 'Harper Lee', 'year': 1960}
}
```

Output:
```
dict_keys(['The Great Gatsby'])
Harper Lee
1960
```

<!-- grade_id: cell-3462e4a5e3da338b, points: 1.5, grade: true, solution: false -->
```python
# Case 9
assert manage_book_records(book_collection, 'remove', 'To Kill a Mockingbird') == "Book 'To Kill a Mockingbird' has been removed."
```

Output:
```
dict_keys(['The Great Gatsby', 'To Kill a Mockingbird'])

0
```

<!-- grade_id: cell-42397c3c546c5ecd, points: 1.5, grade: true, solution: false -->
```python
# Case 10
assert manage_book_records(book_collection, 'remove', 'To Kill a Mockingbird') == "Book 'To Kill a Mockingbird' not found."
```

Output:
```
dict_keys(['The Great Gatsby'])

0
```

---
## (15 คะแนน) <mark>Exercise #8 (Moderate) **การใช้ `list` 2 มิติ**</mark>

- เขียนฟังก์ชัน `transpose_matrix(matrix)` ที่มีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับพารามิเตอร์เป็น `matrix` ซึ่งเป็น `list` 2 มิติ
  2) สร้างเมทริกซ์ใหม่ที่เป็นการ transpose ของเมทริกซ์เดิม (สลับแถวกับคอลัมน์)
  3) return เมทริกซ์ใหม่
- เช่น `transpose_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` จะได้ `[[1, 4, 7], [2, 5, 8], [3, 6, 9]]`

<!-- grade_id: cell-db0c4c209ec24099, points: 0, grade: false, solution: true -->
```python
def transpose_matrix(matrix):
    # YOUR CODE HERE
    m = np.array(matrix)
    m = m.transpose()
    m = m.tolist()
    return m
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-5f9ee0627daa2e97, points: 1.5, grade: true, solution: false -->
```python
# Case 1
assert transpose_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]]) == [[1, 4, 7], [2, 5, 8], [3, 6, 9]]
```

<!-- grade_id: cell-038a41531511686e, points: 1.5, grade: true, solution: false -->
```python
# Case 2
assert transpose_matrix([[1, 2], [3, 4], [5, 6]]) == [[1, 3, 5], [2, 4, 6]]  # Test with 2 columns
```

<!-- grade_id: cell-78ef598bdc5f7474, points: 1.5, grade: true, solution: false -->
```python
# Case 3
assert transpose_matrix([[1]]) == [[1]]  # Test with 1x1 matrix
```

<!-- grade_id: cell-3599c634c5546b44, points: 1.5, grade: true, solution: false -->
```python
# Case 4
assert transpose_matrix([[1, 2, 3]]) == [[1], [2], [3]]  # Test with 1 row matrix
```

<!-- grade_id: cell-965c276fa83e20b9, points: 1.5, grade: true, solution: false -->
```python
# Case 5
assert transpose_matrix([[1, 2], [3, 4]]) == [[1, 3], [2, 4]]  # Test with 2x2 matrix
```

<!-- grade_id: cell-580f64dc2d539cca, points: 1.5, grade: true, solution: false -->
```python
# Case 6
assert transpose_matrix([[1, 2, 3, 4], [5, 6, 7, 8]]) == [[1, 5], [2, 6], [3, 7], [4, 8]]  # Test with 4 columns
```

<!-- grade_id: cell-5052514678bc4610, points: 1.5, grade: true, solution: false -->
```python
# Case 7
assert transpose_matrix([['a', 'b'], ['c', 'd']]) == [['a', 'c'], ['b', 'd']]  # Test with string values
```

<!-- grade_id: cell-e16e7a0093c86da6, points: 1.5, grade: true, solution: false -->
```python
# Case 8
assert transpose_matrix([[1, 2, 3], [4, 5, 6]]) == [[1, 4], [2, 5], [3, 6]]  # Test with 3 rows and 2 columns
```

<!-- grade_id: cell-9ffae084b6b2600d, points: 1.5, grade: true, solution: false -->
```python
# Case 9
assert transpose_matrix([[1, 2], [3, 4], [5, 6], [7, 8]]) == [[1, 3, 5, 7], [2, 4, 6, 8]]  # Test with more rows than columns
```

<!-- grade_id: cell-9bc4b4d6355b40d7, points: 1.5, grade: true, solution: false -->
```python
# Case 10
assert transpose_matrix([[True, False], [False, True]]) == [[True, False], [False, True]]  # Test with boolean values
```

---
## (15 คะแนน) <mark>Exercise #9 (Moderate): **การใช้งาน `set`**</mark>

- เขียนฟังก์ชัน `set_operations(set_a, set_b)` ที่มีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับพารามิเตอร์เป็นเซ็ต `set_a` และ `set_b`
  2) ทำ operation ต่าง ๆ ดังนี้
     - Union
     - Intersection
     - Difference
     - Symmetric Difference
  3) return ผลลัพธ์ในรูปแบบ `dict` โดยมี **key** เป็นชื่อของ operation และ **value** คือ ผลลัพธ์
- เช่น
  - `set_operations({1, 2, 3}, {3, 4, 5})` จะได้ `{'union': {1, 2, 3, 4, 5}, 'intersection': {3}, 'difference': {1, 2}, 'symmetric_difference': {1, 2, 4, 5}}`

<!-- grade_id: cell-61dcedefcbbe6f9d, points: 0, grade: false, solution: true -->
```python
def set_operations(set_a, set_b):
    # YOUR CODE HERE
    set_op = {
    'union': set_a.union(set_b),
    'intersection': set_a.intersection(set_b),
    'difference': set_a.difference(set_b),
    'symmetric_difference': set_a.symmetric_difference(set_b)    
    }
    return set_op
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-94fbf279cf92ec8c, points: 1.5, grade: true, solution: false -->
```python
# Case 1
assert set_operations({1, 2, 3}, {3, 4, 5}) == {
    'union': {1, 2, 3, 4, 5},
    'intersection': {3},
    'difference': {1, 2},
    'symmetric_difference': {1, 2, 4, 5}
}
```

<!-- grade_id: cell-a0394cb445d22540, points: 1.5, grade: true, solution: false -->
```python
# Case 2
assert set_operations({1, 2}, {2, 3}) == {
    'union': {1, 2, 3},
    'intersection': {2},
    'difference': {1},
    'symmetric_difference': {1, 3}
} 
```

<!-- grade_id: cell-fd3536770a2b4cde, points: 1.5, grade: true, solution: false -->
```python
# Case 3
assert set_operations({'a', 'b', 'c'}, {'c', 'd', 'e'}) == {
    'union': {'a', 'b', 'c', 'd', 'e'},
    'intersection': {'c'},
    'difference': {'a', 'b'},
    'symmetric_difference': {'a', 'b', 'd', 'e'}
}
```

<!-- grade_id: cell-a74a14675dede00a, points: 1.5, grade: true, solution: false -->
```python
# Case 4
assert set_operations({1, 2, 3}, {}) == {
    'union': {1, 2, 3},
    'intersection': set(),
    'difference': {1, 2, 3},
    'symmetric_difference': {1, 2, 3}
}
```

<!-- grade_id: cell-65bb6e074548ad89, points: 1.5, grade: true, solution: false -->
```python
# Case 5
assert set_operations({1, 1, 2, 2}, {2, 3}) == {
    'union': {1, 2, 3},
    'intersection': {2},
    'difference': {1},
    'symmetric_difference': {1, 3}
}
```

<!-- grade_id: cell-0cca0b8681aad476, points: 1.5, grade: true, solution: false -->
```python
# Case 6
assert set_operations({1, 2, 3}, {1, 2, 3}) == {
    'union': {1, 2, 3},
    'intersection': {1, 2, 3},
    'difference': set(),
    'symmetric_difference': set()
}
```

<!-- grade_id: cell-87dcc45b1c1b68c8, points: 1.5, grade: true, solution: false -->
```python
# Case 7
assert set_operations({0, 1, 2}, {2, 3, 4}) == {
    'union': {0, 1, 2, 3, 4},
    'intersection': {2},
    'difference': {0, 1},
    'symmetric_difference': {0, 1, 3, 4}
}
```

<!-- grade_id: cell-15f4eca89c368caf, points: 1.5, grade: true, solution: false -->
```python
# Case 8
assert set_operations({10, 20, 30}, {20, 30, 40, 50}) == {
    'union': {10, 20, 30, 40, 50},
    'intersection': {20, 30},
    'difference': {10},
    'symmetric_difference': {10, 40, 50}
}
```

<!-- grade_id: cell-d540e042ba915ed8, points: 1.5, grade: true, solution: false -->
```python
# Case 9
assert set_operations({'apple', 'banana'}, {'banana', 'cherry'}) == {
    'union': {'apple', 'banana', 'cherry'},
    'intersection': {'banana'},
    'difference': {'apple'},
    'symmetric_difference': {'apple', 'cherry'}
}
```

<!-- grade_id: cell-848f9279e4a6cf75, points: 1.5, grade: true, solution: false -->
```python
# Case 10
assert set_operations({1.1, 2.2, 3.3}, {3.3, 4.4}) == {
    'union': {1.1, 2.2, 3.3, 4.4},
    'intersection': {3.3},
    'difference': {1.1, 2.2},
    'symmetric_difference': {1.1, 2.2, 4.4}
}
```

---
## (25 คะแนน) <mark>Exercise #10 (Difficult): **ตรวจเงื่อนไขเกม Sudoku ด้วย NumPy**</mark>

- เขียนฟังก์ชัน `check_sudoku(board)` ที่มีเงื่อนไขดังนี้
  1) ฟังก์ชันนี้จะรับ **NumPy** อาร์เรย์ 2 มิติ (`board`) ขนาด 9x9 ซึ่งเป็นกระดาน Sudoku ซึ่งต้องมีเงื่อนไขดังนี้
      - **การตรวจสอบแถว** แต่ละแถวต้องมีตัวเลข 1-9 โดยไม่มีการซ้ำกัน (ไม่สนใจเลขศูนย์)
      - **การตรวจสอบหลัก** แต่ละหลักต้องมีตัวเลข 1-9 โดยไม่มีการซ้ำกัน (ไม่สนใจเลขศูนย์)
      - **การตรวจสอบกลุ่มย่อย (subgrid)** แต่ละกลุ่มย่อย 3x3 ต้องมีตัวเลข 1-9 โดยไม่มีการซ้ำกัน (ไม่สนใจเลขศูนย์)
        - **subgrid** คือ ตารางย่อย ขนาด 3x3 นับจาก ซ้ายไปขวา และ บนลงล่าง เช่น **subgrid 1** คือ **row: 1-3 และ col: 1-3** หรือ **subgrid 5** คือ **row: 4-6 และ col: 4-6** เป็นต้น
  2) หากผ่านทุกเงื่อนไข ฟังก์ชันจะ return ค่า `'True'`
  3) หากไม่ผ่าน ฟังก์ชันจะ return ข้อผิดพลาดทั้งหมดที่ตรวจพบ ตามเงื่อนไขดังนี้        
      - หากไม่ผ่าน จะ return `()'False',[list_of_errors])` โดย `list` แสดงตามลำดับ คือ `[row_error, col_error, subgrid_error]` มีรายละเอียด ดังนี้
        - ไม่ผ่านเพราะ **การตรวจสอบแถว** return `"Duplicate numbers X in row I"`
        - ไม่ผ่านเพราะ **การตรวจสอบหลัก** return `"Duplicate numbers Y in column J"`
        - ไม่ผ่านเพราะ **การตรวจสอบกลุ่มย่อย** return `"Duplicate numbers Z in subgrid K"`

- เช่น
    ```
    board = np.array([
      [5, 3, 4, 6, 7, 8, 9, 1, 2],
      [6, 7, 2, 1, 9, 5, 3, 4, 8],
      [1, 9, 8, 3, 4, 2, 5, 6, 7],
      [8, 5, 9, 7, 6, 1, 4, 2, 3],
      [4, 2, 6, 8, 5, 3, 7, 9, 1],
      [7, 1, 3, 9, 2, 4, 8, 5, 6],
      [9, 6, 1, 5, 3, 7, 2, 8, 4],
      [2, 8, 7, 4, 1, 9, 6, 3, 5],
      [3, 4, 5, 2, 8, 6, 1, 7, 9]
    ])
    ```
  - ผลลัพธ์ `check_sudoku(board)` จะคืนค่า `'True'`


  ```
  board = np.array([
    [9, 9, 1, 8, 6, 7, 2, 5, 3],
    [8, 2, 6, 3, 9, 5, 1, 7, 4],
    [3, 5, 7, 1, 2, 4, 6, 8, 9],
    [4, 6, 3, 5, 6, 1, 9, 2, 8],
    [7, 1, 2, 6, 8, 9, 4, 3, 5],
    [5, 8, 9, 4, 3, 2, 7, 6, 1],
    [1, 7, 5, 2, 4, 8, 3, 9, 6],
    [2, 3, 4, 9, 5, 6, 8, 1, 7],
    [6, 9, 8, 7, 1, 3, 5, 4, 2]
  ])
  ```
  -  ผลลัพธ์ `check_sudoku(board)` จะคืนค่า `('False', ['Duplicate numbers 9 in row 1', 'Duplicate numbers 6 in row 4', 'Duplicate numbers 9 in column 2', 'Duplicate numbers 6 in column 5', 'Duplicate numbers 9 in subgrid 1', 'Duplicate numbers 6 in subgrid 5'])`

<!-- grade_id: cell-ef97922ae6f0e529, points: 0, grade: false, solution: true -->
```python
def check_sudoku(board):
    # YOUR CODE HERE
    row_error = []
    col_error = []
    sub_error = []
    li_subgrid = []
    for row in range(0,7,3):
        for col in range(0,7,3):
            li_subgrid.append(board[row:row+3,col:col+3])
        """ subgrid = [board[:3,:3]
             ,board[:3,3:6]
             ,board[:3,6:9]
             ,board[3:6,:3]
             ,board[3:6,3:6]
             ,board[3:6,6:9]
             ,board[6:9,:3]
             ,board[6:9,3:6]
             ,board[6:9,6:9]] """
    for i in range(9):
        num_row , count_row =np.unique(board[i:i+1,:], return_counts=True)
        zip_row = dict(zip(num_row,count_row))
        head_row = list(zip_row.keys())
        c_row = list(zip_row.values())
        for j in c_row:#check row
            if j > 1:
                row_error.append(f"Duplicate numbers {head_row[c_row.index(j)]} in row {i+1}")
        #print(f'row {i}',counter_row)
        num_cu , count_cu = np.unique(board[:,i:i+1], return_counts=True)
        zip_cu = dict(zip(num_cu,count_cu))
        head_cu = list(zip_cu.keys())
        c_cu = list(zip_cu.values())
        for k in c_cu:# column
            if k > 1:
                col_error.append(f"Duplicate numbers {head_cu[c_cu.index(k)]} in column {i+1}")
        #print(f'cul {i}',counter_cu)
    for i in range(len(li_subgrid)):
        c = 0
        num_sub , count_sub = np.unique(li_subgrid[i],return_counts=True)
        zip_sub = dict(zip(num_sub,count_sub))
        head_sub = list(zip_sub.keys())
        c_sub = list(zip_sub.values())
        for j in c_sub:# li_subgrid
            if j >1:
                sub_error.append(f"Duplicate numbers {head_sub[c_sub.index(j)]} in subgrid {i+1}")
    ans = row_error+col_error+sub_error
    if len(ans) != 0:
        print(ans)
        return 'False',ans
    else:
        return 'True'
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<!-- grade_id: cell-f23c9f23d2024a0d, points: 5, grade: true, solution: false -->
```python
# Case 1
board1 = np.array([
    [8, 3, 5, 4, 1, 6, 9, 2, 7],
    [2, 9, 6, 8, 5, 7, 4, 3, 1],
    [4, 1, 7, 9, 2, 3, 6, 5, 8],
    [5, 6, 9, 1, 3, 4, 7, 8, 2],
    [1, 2, 3, 6, 7, 8, 5, 9, 4],
    [7, 4, 8, 5, 9, 2, 1, 6, 3],
    [6, 5, 2, 7, 8, 1, 3, 4, 9],
    [9, 8, 1, 3, 4, 5, 2, 7, 6],
    [3, 7, 4, 2, 6, 9, 8, 1, 5]
])
assert check_sudoku(board1) == 'True'
```

<!-- grade_id: cell-5a57a7c9cc5cb1f3, points: 5, grade: true, solution: false -->
```python
# Case 2
board2 = np.array([
    [8, 3, 6, 2, 5, 7, 9, 4, 1],
    [2, 1, 9, 3, 6, 4, 8, 7, 5],
    [5, 4, 7, 1, 9, 8, 2, 6, 3],
    [4, 7, 5, 6, 2, 1, 3, 2, 9],
    [9, 6, 2, 5, 7, 3, 4, 1, 8],
    [3, 8, 1, 4, 2, 9, 6, 5, 7],
    [6, 2, 8, 9, 1, 5, 7, 3, 4],
    [1, 9, 3, 7, 4, 6, 5, 8, 2],
    [7, 5, 4, 8, 3, 2, 1, 9, 6]
])
assert check_sudoku(board2) == ('False', ['Duplicate numbers 2 in row 4', 'Duplicate numbers 2 in column 5', 'Duplicate numbers 2 in subgrid 5'])
```

Output:
```
['Duplicate numbers 2 in row 4', 'Duplicate numbers 2 in column 5', 'Duplicate numbers 2 in subgrid 5']
```

<!-- grade_id: cell-046d221069ed3da4, points: 5, grade: true, solution: false -->
```python
# Case 3
board3 = np.array([
    [5, 4, 7, 6, 7, 2, 3, 8, 1],
    [2, 6, 9, 3, 8, 1, 7, 4, 5],
    [3, 1, 8, 7, 4, 5, 6, 9, 2],
    [4, 2, 5, 1, 3, 6, 9, 7, 8],
    [8, 9, 3, 2, 7, 4, 1, 5, 6],
    [6, 7, 1, 9, 5, 8, 2, 3, 4],
    [1, 3, 4, 8, 2, 7, 5, 6, 9],
    [9, 8, 6, 5, 1, 3, 4, 2, 7],
    [7, 5, 2, 4, 6, 9, 8, 1, 3]
])

assert check_sudoku(board3) == ('False', ['Duplicate numbers 7 in row 1', 'Duplicate numbers 7 in column 5', 'Duplicate numbers 7 in subgrid 2'])
```

Output:
```
['Duplicate numbers 7 in row 1', 'Duplicate numbers 7 in column 5', 'Duplicate numbers 7 in subgrid 2']
```

<!-- grade_id: cell-dbd02426bfb27b13, points: 5, grade: true, solution: false -->
```python
# Case 4
board4 = np.array([
    [8, 3, 5, 4, 1, 6, 9, 2, 7],
    [2, 9, 6, 8, 5, 7, 4, 3, 1],
    [4, 2, 7, 9, 2, 3, 2, 5, 8],
    [5, 6, 9, 1, 3, 4, 7, 8, 2],
    [1, 2, 3, 6, 7, 8, 5, 9, 4],
    [7, 4, 8, 5, 9, 2, 1, 6, 3],
    [6, 5, 2, 7, 8, 1, 3, 4, 9],
    [9, 8, 1, 3, 4, 5, 2, 7, 6],
    [3, 7, 4, 2, 6, 9, 8, 1, 5]
])

assert check_sudoku(board4) == ('False', ['Duplicate numbers 2 in row 3', 'Duplicate numbers 2 in column 2', 'Duplicate numbers 2 in column 7', 'Duplicate numbers 2 in subgrid 1', 'Duplicate numbers 2 in subgrid 3'])
```

Output:
```
['Duplicate numbers 2 in row 3', 'Duplicate numbers 2 in column 2', 'Duplicate numbers 2 in column 7', 'Duplicate numbers 2 in subgrid 1', 'Duplicate numbers 2 in subgrid 3']
```

<!-- grade_id: cell-521a63ba8ccbab3e, points: 5, grade: true, solution: false -->
```python
# Case 5
board5 = np.array([
    [9, 9, 1, 8, 6, 7, 2, 5, 3],
    [8, 2, 6, 3, 9, 5, 1, 7, 4],
    [3, 5, 7, 1, 2, 4, 6, 8, 9],
    [4, 6, 3, 5, 6, 1, 9, 2, 8],
    [7, 1, 2, 6, 8, 9, 4, 3, 5],
    [5, 8, 9, 4, 3, 2, 7, 6, 1],
    [1, 7, 5, 2, 4, 8, 3, 9, 6],
    [2, 3, 4, 9, 5, 6, 8, 1, 7],
    [6, 9, 8, 7, 1, 3, 5, 4, 2]
])

assert check_sudoku(board5) == ('False', ['Duplicate numbers 9 in row 1', 'Duplicate numbers 6 in row 4', 'Duplicate numbers 9 in column 2', 'Duplicate numbers 6 in column 5', 'Duplicate numbers 9 in subgrid 1', 'Duplicate numbers 6 in subgrid 5'])
```

Output:
```
['Duplicate numbers 9 in row 1', 'Duplicate numbers 6 in row 4', 'Duplicate numbers 9 in column 2', 'Duplicate numbers 6 in column 5', 'Duplicate numbers 9 in subgrid 1', 'Duplicate numbers 6 in subgrid 5']
```