Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# SECTION 2 ONLY
Name = ""

Student ID = ""

# PART-1: String - Files and Exception

**Write you codes after `# YOUR CODE HERE`**

Use the **provided files**

<mark>**DO NOT IMPORT EXTERNAL LIBRARY**</mark>

<mark>**DO NOT CHANGE CONTENT OF TEXT AND CSV FILES**</mark>

---
## <mark>Exercise #1 
**Finds word in file, returns line number** (Total: 5 points) </mark>
- Comple a `find_word()` function that reads a text file to:
    - Searches for a specific word in a text file.
    - returns the line number where it was found.

- For example:
    - `find_word('filename.txt', 'r', 'Sydney')` returns `1`, indicating that `Sydney` was found on line 1.
    - `find_word('filename.txt', 'r', 'istanbul')` returns `-1`, indicating that `istanbul` was not found.


<!-- grade_id: cell-d73ce7ef779cd16a, points: 0, grade: false, solution: true -->
```python
file_name = 'words.txt'
read_mode = 'r'

def find_word(file_name, read_mode, search_word):
    c = 0
    # YOUR CODE HERE
    with open(file_name,read_mode) as f:
        for i in f:
            c += 1
            check = i.strip()
            if  check.lower() == search_word.lower():
                return c
        return -1
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **5 points**

<!-- grade_id: cell-755b8cc5ab24d5ff, points: 1, grade: true, solution: false -->
```python
# 1 point
assert(find_word(file_name, read_mode, 'Paris') == 2)
```

<!-- grade_id: cell-e1d3dbfea96f642e, points: 1, grade: true, solution: false -->
```python
# 1 point
assert(find_word(file_name, read_mode, 'Shanghai') == 6)
```

<!-- grade_id: cell-cce741d6d22f63a7, points: 1, grade: true, solution: false -->
```python
# 1 point
assert(find_word(file_name, read_mode, 'Cairo') == 9)
```

<!-- grade_id: cell-390f2025b1895e7d, points: 1, grade: true, solution: false -->
```python
# 1 point
assert(find_word(file_name, read_mode, 'new york') == 4)
```

<!-- grade_id: cell-4bbac234b1752d45, points: 1, grade: true, solution: false -->
```python
# 1 point
assert(find_word(file_name, read_mode, 'Rio de Janeiro') == -1)
```

---
## <mark>Exercise #2 
**Finds number in file, returns line number** (Total: 5 points) </mark>
- Comple a `find_number()` function that reads a text file to:
    - Searches for a specific number in a text file.
    - returns the line number where it was found.

- For example:
    - `find_number('filename.txt', 'r', 10)` returns `1`, indicating that `10` was found on line `1`.
    - `find_number('filename.txt', 'r', 999)` returns `-1`, indicating that `999` was not found.

<!-- grade_id: cell-f4a80c0b38839e0b, points: 0, grade: false, solution: true -->
```python
file_name = 'numbers.txt'
read_mode = 'r'

def find_number(file_name, read_mode, search_number):

    # YOUR CODE HERE
    c = 0
    with open(file_name,read_mode) as f:
        for i in f:
            check = i.strip()
            c += 1
            if  int(check) == search_number:
                return c
        return -1
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **5 points**

<!-- grade_id: cell-c322a97358cde8ae, points: 1, grade: true, solution: false -->
```python
# 1 points
assert(find_number(file_name, read_mode, 4) == 1)
```

<!-- grade_id: cell-ec032c8f6ce066ca, points: 1, grade: true, solution: false -->
```python
# 1 points
assert(find_number(file_name, read_mode, 2) == 3)
```

<!-- grade_id: cell-6604d1f0de0f7768, points: 1, grade: true, solution: false -->
```python
# 1 points
assert(find_number(file_name, read_mode, 12) == 4)
```

<!-- grade_id: cell-d7b98a68a9b35ffd, points: 1, grade: true, solution: false -->
```python
# 1 points
assert(find_number(file_name, read_mode, 9) == 5)
```

<!-- grade_id: cell-64cd0bbcbb3d2660, points: 1, grade: true, solution: false -->
```python
# 1 points
assert(find_number(file_name, read_mode, 1) == -1)
```

---
## <mark>Exercise #3 
**Count Words in a Text File** (Total: 30 points) </mark>
- Create a `read_and_count()` function that reads a text file to:
    - Count the words from the text file.
    - Return the total number of words.


- NOTE: You may use `.split()` to split word from sentence
    - For example, 

        `txt = 'Good Morning'`

        `txt2 = txt.split()`
        
        `print(txt2)`, you will get `['Good', 'Morning']`

- For example:
    - `read_and_count('song.txt', 'green')` returns `1`, indicating that `green` occurred only once.


<!-- grade_id: cell-4db7be5c9fcfb010, points: 0, grade: false, solution: true -->
```python
def read_and_count(file, word):
    # YOUR CODE HERE
    c = 0
    with open(file, 'r') as song:
        for text in song:
            li = text.lower().split()
            c += li.count(word)
        return c
```

```python
# TEST
print(read_and_count('song.txt', 'green'))  # You will get: 1
print(read_and_count('song.txt', 'you'))    # You will get: 3
```

Output:
```
1
3
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **30 points**

<!-- grade_id: cell-a93c1843b6e09c32, points: 7.5, grade: true, solution: false -->
```python
# 7.5 points
assert(read_and_count('song.txt', 'of') == 5)
```

<!-- grade_id: cell-3b69e9135c90a845, points: 7.5, grade: true, solution: false -->
```python
# 7.5 points
assert(read_and_count('song.txt', 'world') == 4)
```

<!-- grade_id: cell-15a44fd437c40429, points: 7.5, grade: true, solution: false -->
```python
# 7.5 points
assert(read_and_count('song.txt', 'i') == 11)
```

<!-- grade_id: cell-5867bf2141f75299, points: 7.5, grade: true, solution: false -->
```python
# 7.5 points
assert(read_and_count('song.txt', 'the') == 6)
```

---
## <mark>Exercise #4 
**Handles programming errors with try-except:** (Total: 5 points) </mark>
- Complete a `perform_operations(file_name, n)` function that reads a text file to:
    - Read the first line in the given file.
    - Return the result of the division of `first line / n`, or:
    - Return `E01` as an error code if an invalid value is provided for its argument.
    - Return `E02` as an error code if division by zero occurs.    
    - Return `E03` as an error code if the file is not found.

- For example:
    - `perform_operations('filename.txt', 5)` returns `2`, indicating that `10` was found in the first line and divided by `5` equals `2`.
    - `perform_operations('no_file.txt', 10)` returns `E03`, indicating that `no_file.txt` was not found.

<!-- grade_id: cell-f2799c2fb8cf0251, points: 0, grade: false, solution: true -->
```python
def perform_operations(file_name, n):
    # YOUR CODE HERE
    try:
        with open(file_name, 'r') as number:
            try:
                res = int(number.readline())/n
                return res
            except ZeroDivisionError:
                return 'E02'
            except ValueError:
                return 'E01'
    except:
        return "E03"
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **10 points**

<!-- grade_id: cell-35c8ae6a58a12a0c, points: 2, grade: true, solution: false -->
```python
# 2 points
assert(perform_operations('numbers.txt', 2) == 2.0)
```

<!-- grade_id: cell-17b1d3948eb00a4f, points: 2, grade: true, solution: false -->
```python
# 2 points
assert(perform_operations('numbers.txt', 5) == 0.8)
```

<!-- grade_id: cell-a26b2e25ca62ae94, points: 2, grade: true, solution: false -->
```python
# 2 points
assert(perform_operations('words.txt', 10) == 'E01')
```

<!-- grade_id: cell-2d02aa42dc598b3a, points: 2, grade: true, solution: false -->
```python
# 2 points
assert(perform_operations('numbers.txt', 0) == 'E02')
```

<!-- grade_id: cell-fdef6af197a5d16f, points: 2, grade: true, solution: false -->
```python
# 2 points
assert(perform_operations('no_file.txt', 10) == 'E03')
```