Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# WEEK 05: Repetition Structure

**Name**: Thanakrit Samoena

**ID**: 67070060

**Program**: AIT

**Intructions: Write you codes after `# YOUR CODE HERE`**

**Automatic grading, DO NOT CHANGE anything outside `# YOUR CODE HERE`**

---
<span style='background :yellow' > Execise #1 </span>.
Create a program to calculate the **sum** of numbers from `1` to `N` using a `for`-loop. **(20 Points)**

**Variables:**
- `N` is a variable to store an integer number as the last number.
- `sum_numbers` is a variable to store the sum of all numbers.

Examples:

- If `N` = 5, then `sum_numbers` of `1:5` = `15` # 1 + 2 + 3 + 4 + 5
- If `N` = 7, then `sum_numbers` of `1:7` = `28` # 1 + 2 + 3 + 4 + 5 + 6 + 7

<!-- grade_id: cell-c2c51a3819bb15f7, points: 0, grade: false, solution: true -->
```python
# Test-Ex1-1 condition

N = 10
sum_numbers = 0

for num in range(1, N + 1):
    
    # YOUR CODE HERE
    sum_numbers += num
print(sum_numbers)
```

Output:
```
55
```

<!-- grade_id: cell-aea3c70ce6e8925c, points: 10, grade: true, solution: false -->
```python
# Test-Ex1-1: a correct solution gets 10 points
assert sum_numbers == 55
```

<!-- grade_id: cell-31e0a36c7290bb7b, points: 0, grade: true, solution: true -->
```python
# Test-Ex1-2 condition

N = 1000
sum_numbers = 0

for num in range(1, N + 1):

    # YOUR CODE HERE
    sum_numbers += num
print(sum_numbers)
```

Output:
```
500500
```

<!-- grade_id: cell-473547942b45caf6, points: 10, grade: true, solution: false -->
```python
# Test-Ex1-2: a correct solution gets 10 points
assert sum_numbers == 500500
```

---
<span style='background :yellow' > Execise #2 </span>.
Create a program to sum **odd** numbers from `1` to `N` using `for`-loop. **(20 Points)**

*Hint: You may need to use* `if` *statement.*

**Variables:**
- `N` is a variable to store an integer number as a last number
- `sum_even` is a variable to store summation of all even numbers

**Examples:**

- If `N` = 5, then the `sum_odd` of `1:5` = `9`    # 1 + 3 + 5
- If `N` = 10, then the `sum_odd` of `1:10` = `25`    # 1 + 3 + 5 + 7 + 9

<!-- grade_id: cell-bb7daf38499e0fa2, points: 0, grade: false, solution: true -->
```python
# Test-Ex2-1 condition

N = 27
sum_odd = 0

for num in range(1, N + 1):

# YOUR CODE HERE
    if num % 2 == 1:
        sum_odd += num
print(sum_odd)

```

Output:
```
196
```

<!-- grade_id: cell-b5bc05487d22a580, points: 10, grade: true, solution: false -->
```python
# Test-Ex2-1: a correct solution gets 10 points
assert sum_odd == 196
```

<!-- grade_id: cell-592f97713a7768c6, points: 0, grade: false, solution: true -->
```python
# Test-Ex2-2 condition

N = 49
sum_odd = 0

for num in range(1, N + 1):
    
# YOUR CODE HERE
    if num % 2 == 1:
        sum_odd += num
print(sum_odd)

```

Output:
```
625
```

<!-- grade_id: cell-d454b0509c7c3055, points: 10, grade: true, solution: false -->
```python
# Test-Ex2-2: a correct solution gets 10 points
assert sum_odd == 625
```

---
<span style='background :yellow' > Execise #3 </span>.
Create a program to sum **even** numbers from `1` to `N` using `while`-loop. **(20 Points)**

*Hint: You may need to use* `if` *statement.*

**Variables:**
- `N` is a variable to store an integer number as a last number
- `sum_even` is a variable to store summation of all even numbers
- `num` is a variable to set a pretest loop of while

**Examples:**

- If `N` = 5, then the `sum_even` of `1:5` = `6`         # 2 + 4
- If `N` = 10, then the `sum_even` of `1:10` = `30`      # 2 + 4 + 6 + 8 + 10

<!-- grade_id: cell-4195fb0f06de1f28, points: 0, grade: false, solution: true -->
```python
# Test-Ex3-1 condition

N = 23
sum_even = 0

num = 1 # pretest loop of while
while num <= N:

# YOUR CODE HERE
    if not num % 2 :
        sum_even += num
    num += 1
print(sum_even)

```

Output:
```
132
```

<!-- grade_id: cell-a0d0d2921f2b4f2e, points: 10, grade: true, solution: false -->
```python
# Test-Ex3-1: a correct solution gets 10 points
assert sum_even == 132
```

<!-- grade_id: cell-37c1d4cba1a5877b, points: 0, grade: false, solution: true -->
```python
# Test-Ex3-2 condition

N = 598
sum_even = 0

num = 1 # pretest loop of while
while num <= N:

# YOUR CODE HERE
    if not num % 2 :
        sum_even += num
    num += 1
print(sum_even)

```

Output:
```
89700
```

<!-- grade_id: cell-1823cc9d93652768, points: 10, grade: true, solution: false -->
```python
# Test-Ex3-2: a correct solution gets 10 points
assert sum_even == 89700
```

---
<span style='background :yellow' > Execise #4 </span>.
Create a program to find the **minimum** number in a list of numbers using a `for`-loop or a `while`-loop. **(20 Points)**

Hint: Initialize a `variable` that keeps a **reference** with a **very large** number.

**Variables:**

- `numbers` is a list of numbers.
- `min_num` is a variable to store the **minimum** number.

**Examples:**

- If `numbers` = [3, 7, 2, 8, 5], then `min_num` = `2`
- If `numbers` = [1, 4, 9, 2, 6], then `min_num` = `1`

<!-- grade_id: cell-e615b04c02aa810e, points: 0, grade: false, solution: true -->
```python
# Test-Ex4-1 condition

numbers = [27, -83, 45, 12, -98, 57, 34, 76, 89, 23]

# YOUR CODE HERE
min_num = numbers[0]
for i in numbers:
    if i < min_num:
        min_num = i
print(min_num)
```

Output:
```
-98
```

<!-- grade_id: cell-71e2f605267a29c8, points: 10, grade: true, solution: false -->
```python
# Test-Ex4-1: a correct solution gets 10 points
assert min_num == -98
```

<!-- grade_id: cell-b1692e55a1901b9c, points: 0, grade: false, solution: true -->
```python
# Test-Ex4-2 condition

numbers = [-23.45, 78.92, -54.67, 33.21, -12.34, 88.45, -67.89, 45.67, -90.12, 22.33]

# YOUR CODE HERE
min_num = numbers[0]
for i in numbers:
    if i < min_num:
        min_num = i
print(min_num)

```

Output:
```
-90.12
```

<!-- grade_id: cell-6e93c1255d469aa8, points: 10, grade: true, solution: false -->
```python
# Test-Ex4-2: a correct solution gets 10 points
assert min_num == -90.12
```

---
<span style='background :yellow' > Execise #5 </span>.
Create a program to count the number in the `secret` code that is greater than a `ref_num` using a `for`-loop **or** a `while`-loop. **(20 Points)**

*Hint: You may need to use* `if` *statement.*

**Variables:**
- `secret` is a list of 12-numbers
- `ref_num` is a reference number
- `count_num` is a variable to keep count of numbers that have a value greater than a `ref_num`

**Examples:**

- If `ref_num` = -13, then the `count_num` is `12`
- If `ref_num` = -10, then the `count_num` is `11`

<!-- grade_id: cell-655163fc08b5cf6a, points: 0, grade: false, solution: true -->
```python
# Test-Ex5-1 condition

secret = [10, 20, 33, 21, 34, 7, 9, 20, 99, -12, 3,87]

ref_num = 23
count_num = 0

# YOUR CODE HERE
for i in secret:
    if ref_num < i:
        count_num += 1
print(count_num)

```

Output:
```
4
```

<!-- grade_id: cell-f4e5d2645a541123, points: 10, grade: true, solution: false -->
```python
# Test-Ex5-1: a correct solution gets 10 points
assert count_num == 4
```

<!-- grade_id: cell-e8ff021a68beb143, points: 0, grade: false, solution: true -->
```python
# Test-Ex5-2 condition

secret = [10, 20, 33, 21, 34, 7, 9, 20, 99, -12, 3,87]

ref_num = 10
count_num = 0

# YOUR CODE HERE
for i in secret:
    if ref_num < i:
        count_num += 1
print(count_num)


```

Output:
```
7
```

<!-- grade_id: cell-6bde0f4b8578b263, points: 10, grade: true, solution: false -->
```python
# Test-Ex5-2: a correct solution gets 10 points
assert count_num == 7
```