Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# SECTION 2 ONLY
Name = "Thanakrit Samoena"

Student ID = "67070060"

# String - Exercise

**Write you codes after `# YOUR CODE HERE`**

*ใช้ **ASCII.pdf** เป็นข้อมูลอ้างอิง (ถ้าจำเป็น)*

<mark>**DO NOT IMPORT EXTERNAL LIBRARY OTHER THAN `string` **</mark>

```python
import string   # Use for all problems (if any)
```

---
## <mark>Exercise #1 **Count Consonants** (Total: 5 points)</mark>
- พยัญชนะ (Consonants) คือ ตัวอักษรทั้งหมดในอักษรภาษาอังกฤษที่ไม่ใช่สระ
- สร้างฟังก์ชัน `count_consonants()` ที่มีเงื่อนไขดังนี้
    1) นับจำนวนพยัญชนะทั้งหมดในสตริงที่กำหนด
    2) Return สตริงใหม่ที่ประกอบด้วยพยัญชนะ (ไม่มีสระ) จากสตริงต้นฉบับ
- ตัวอย่างเช่น:
    - `count_consonants('Hello World!')` จะได้เป็น `'HllWrld'` และจำนวนพยัญชนะทั้งหมดคือ `7`

<!-- grade_id: cell-e1f0417099ed9857, points: 0, grade: false, solution: true -->
```python
def count_consonants(input_string):
    # YOUR CODE HERE
    con = "AEIOUaieou"
    res = ""
    count_con = 0
    for i in input_string:
        if i not in con and i.isalpha():
            count_con += 1
            res += i
    return res,count_con
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**0.5 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-b6e6c1d11e23694b, points: 0.5, grade: true, solution: false -->
```python
# Test Case 1
assert count_consonants('Hello World!') == ('HllWrld', 7)
```

<!-- grade_id: cell-18e342be46315529, points: 0.5, grade: true, solution: false -->
```python
# Test Case 2
assert count_consonants('Python Programming') == ('PythnPrgrmmng', 13)
```

<!-- grade_id: cell-07d32bf132bdb9c4, points: 0.5, grade: true, solution: false -->
```python
# Test Case 3
assert count_consonants('a e i o u') == ('', 0)
```

<!-- grade_id: cell-9424820050e136f3, points: 0.5, grade: true, solution: false -->
```python
# Test Case 4
```

<!-- grade_id: cell-dc99a17b589c20ce, points: 0.5, grade: true, solution: false -->
```python
# Test Case 5
```

<!-- grade_id: cell-15ce65039e3b79ce, points: 0.5, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-5efa13f15f2db50a, points: 0.5, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-c39030f1dc76b5d4, points: 0.5, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-d3db83b7928e7f13, points: 0.5, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-86868e2afff4c9a8, points: 0.5, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #2 **Reverse Words** (Total: 5 points)</mark>
- สร้างฟังก์ชัน `reverse_words()` ที่รับประโยคเป็นพารามิเตอร์ตามเงื่อนไขดังนี้
    1) กลับลำดับของคำในประโยค
    2) Return จำนวนคำทั้งหมดในประโยค
- ตัวอย่างเช่น:
    - `reverse_words('Hello World from Python')` จะได้ `'Python from World Hello'` และจำนวนคำทั้งหมดคือ `4`

<!-- grade_id: cell-ed46ce1e8f914481, points: 0, grade: false, solution: true -->
```python
def reverse_words(input_string):
    # YOUR CODE HERE
    return ' '.join(input_string.split()[::-1]),len(input_string.split())
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**0.5 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-f6c9681d601ae7d0, points: 0.5, grade: true, solution: false -->
```python
# Test Case 1
assert reverse_words('Hello World from Python') == ('Python from World Hello', 4)
```

<!-- grade_id: cell-ad23fda3af62780d, points: 0.5, grade: true, solution: false -->
```python
# Test Case 2
assert reverse_words('This is a test') == ('test a is This', 4)
```

<!-- grade_id: cell-62639f8e51da309f, points: 0.5, grade: true, solution: false -->
```python
# Test Case 3
assert reverse_words('I love coding') == ('coding love I', 3)
```

<!-- grade_id: cell-da43882dc9843b91, points: 0.5, grade: true, solution: false -->
```python
# Test Case 4
assert reverse_words('One Two Three Four') == ('Four Three Two One', 4)
```

<!-- grade_id: cell-7551e0ce06bc60f8, points: 0.5, grade: true, solution: false -->
```python
# Test Case 5
assert reverse_words('') == ('', 0)
```

<!-- grade_id: cell-c773293861e026ea, points: 0.5, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-4c88bd110f0a4fd2, points: 0.5, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-7cd4956229622b96, points: 0.5, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-646023e20cc7a1c3, points: 0.5, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-8ddeeb3b310db617, points: 0.5, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #3 **Character Frequency** (Total: 5 points)</mark>
- สร้างฟังก์ชัน `char_frequency()` ที่รับสตริงเป็นพารามิเตอร์และ return คำตอบในรูปแบบ `dictionary` ที่แสดงความถี่ที่ของแต่ละอักขระในสตริง
- ให้ละเว้นช่องว่างและนับอักขระตัวพิมพ์ใหญ่และตัวพิมพ์เล็กเป็นตัวเดียวกัน
- ตัวอย่างเช่น:
    - `char_frequency('Hello World')` จะต้องส่งคืน `{'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1}`

<!-- grade_id: cell-f093826fc912c56e, points: 0, grade: false, solution: true -->
```python
def char_frequency(input_string):
    # YOUR CODE HERE
    txt = input_string.lower()
    cha_freq = {}
    for i in txt:
        if not i.isspace():
            cha_freq[i] = txt.count(i)
    return cha_freq
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**0.5 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-21c1a6423b9eb03d, points: 0.5, grade: true, solution: false -->
```python
# Test Case 1
assert char_frequency('Hello World') == {'h': 1, 'e': 1, 'l': 3, 'o': 2, 'w': 1, 'r': 1, 'd': 1}
```

<!-- grade_id: cell-2a466f4324a36513, points: 0.5, grade: true, solution: false -->
```python
# Test Case 2
assert char_frequency('Python Programming') == {'p': 2, 'y': 1, 't': 1, 'h': 1, 'o': 2, 'n': 2, 'r': 2, 'g': 2, 'a':1, 'm': 2, 'i': 1}
```

<!-- grade_id: cell-cc36e066f4734dd7, points: 0.5, grade: true, solution: false -->
```python
# Test Case 3
assert char_frequency('A quick brown fox!') == {'a': 1, 'q': 1, 'u': 1, 'i': 1, 'c': 1, 'k': 1, 'b': 1, 'r': 1, 'o': 2, 'w': 1, 'n': 1, 'f': 1, 'x': 1, '!': 1}
```

<!-- grade_id: cell-7c2d0f23febeef51, points: 0.5, grade: true, solution: false -->
```python
# Test Case 4
assert char_frequency('12345') == {'1': 1, '2': 1, '3': 1, '4': 1, '5': 1}
```

<!-- grade_id: cell-0e0479598c305671, points: 0.5, grade: true, solution: false -->
```python
# Test Case 5
assert char_frequency('!!@@##$$%%^^&&') == {'!': 2, '@': 2, '#': 2, '$': 2, '%': 2, '^': 2, '&': 2}
```

<!-- grade_id: cell-603da45fb187bd18, points: 0.5, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-1dd58f30876eb598, points: 0.5, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-e191d5b62953a039, points: 0.5, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-ae9c113ca34e78a8, points: 0.5, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-df8819481cb1bd6e, points: 0.5, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #4 **Remove Punctuation** (Total: 5 points)</mark>
- สร้างฟังก์ชัน `remove_punctuation()` ที่รับสตริงเป็นพารามิเตอร์และลบอักขระที่เป็นเครื่องหมายวรรคตอน (Punctuation) ทั้งหมดออก
- ฟังก์ชัน return สตริงที่ลบอักขระที่เป็นเครื่องหมายวรรคตอนออกแล้ว
- ตัวอย่างเช่น:
    - `remove_punctuation('Hello, World! How are you?')` จะได้ `'Hello World How are you'`

- หมายเหตุ อักขระที่ถือว่าเป็นเครื่องหมายวรรคตอน คือ `., ?, !, :, ;, ', ", (, ), [, ], {, }, -, –, ..., /`

<!-- grade_id: cell-6a78983667cdb6f4, points: 0, grade: false, solution: true -->
```python
def remove_punctuation(input_string):
    # YOUR CODE HERE
    c =""
    for i in input_string:
        if i.isalpha() or i.isspace():
            c += i
    return c

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**0.5 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-8d40208f34e546fb, points: 0.5, grade: true, solution: false -->
```python
# Test Case 1
assert remove_punctuation('Hello, World! How are you?') == 'Hello World How are you'
```

<!-- grade_id: cell-02cf9e51879d5f69, points: 0.5, grade: true, solution: false -->
```python
# Test Case 2
assert remove_punctuation('Python is fun; let\'s code!') == 'Python is fun lets code'
```

<!-- grade_id: cell-c88226a78da66cc6, points: 0.5, grade: true, solution: false -->
```python
# Test Case 3
assert remove_punctuation('') == ''
```

<!-- grade_id: cell-1bacf9e0deb0f11b, points: 0.5, grade: true, solution: false -->
```python
# Test Case 4
assert remove_punctuation('No punctuation here') == 'No punctuation here'
```

<!-- grade_id: cell-0126b07bca2bd6d8, points: 0.5, grade: true, solution: false -->
```python
# Test Case 5
assert remove_punctuation('C++ is a programming language.') == 'C is a programming language'
```

<!-- grade_id: cell-35901cf17f59cd96, points: 0.5, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-5936328260544f25, points: 0.5, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-8d957160d9abaca2, points: 0.5, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-ffde811ac92cd962, points: 0.5, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-a861b7c32705bd8d, points: 0.5, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #5 **Vowel to Number** (Total: 10 points) </mark>

สร้างฟังก์ชัน `vowel2num()` ที่แปลงสตริงที่กำหนดโดยการแทนที่สระด้วยตัวเลขเฉพาะและนับจำนวนสระทั้งหมดในสตริงนั้น ตามเงื่อนไข ดังนี้
1. **ระบุสระ:**
   - สระที่ต้องแทนที่คือ `a`, `e`, `i`, `o`, และ `u` (**ทั้งตัวพิมพ์เล็กและตัวพิมพ์ใหญ่**)

2. **ข้อกำหนดของฟังก์ชัน:**
   - ฟังก์ชัน `vowel2num()` จะแทนที่ตัวอักษรด้วยสระ ดังนี้
     1. แทนที่สระแต่ละตัวในสตริงด้วยตัวเลขที่กำหนด
        - `A` หรือ `a` → `1`
        - `E` หรือ `e` → `2`
        - `I` หรือ `i` → `3`
        - `O` หรือ `o` → `4`
        - `U` หรือ `u` → `5`
     2. นับและ return จำนวนรวมของสระในสตริง

3. **คำตอบของฟังก์ชัน**
   - ฟังก์ชันจะต้อง return เป็น `tuple` ที่ประกอบด้วย
     - สตริงใหม่ที่มีสระแทนที่ด้วยตัวเลขที่กำหนด
     - จำนวนเต็มที่แสดงถึงจำนวนสระที่ถูกแทนที่ด้วยตัวเลข

- ตัวอย่าง: `vowel2num('a-e-i-o-u')` จะได้ `('1-2-3-4-5', 5)` หมายถึง สตริงไหม่คือ `'1-2-3-4-5'` มีจำนวนสระที่ถูกแทนด้วยตัวเลขทั้งหมด `5` ตำแหน่ง
- หมายเหตุ:ไม่แยกแยะตัวพิมพ์ใหญ่และพิมพ์เล็ก เช่น `A` และ `a` จะถือว่าเป็นสระตัวเดียวกัน

<!-- grade_id: cell-3d8466b6b5aee272, points: 0, grade: false, solution: true -->
```python
def vowel2num(words):
    # YOUR CODE HERE
    con = "AEIOU"
    res = ""
    c = 0
    for i in words:
        if i.upper() in con:
            res += str(con.index(i.upper())+1)
            c += 1
        else:
            res += i            
    return res,c
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**1 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-1a2797c4f6c4d298, points: 1, grade: true, solution: false -->
```python
# Test Case 1
assert(vowel2num('one cat eats a lot of fishes') == ('4n2 c1t 21ts 1 l4t 4f f3sh2s', 10))
```

<!-- grade_id: cell-6321798011c31ae7, points: 1, grade: true, solution: false -->
```python
# Test Case 2
assert(vowel2num('the queen likes coffee and toast') == ('th2 q522n l3k2s c4ff22 1nd t41st', 12))
```

<!-- grade_id: cell-3742e6e7524b710d, points: 1, grade: true, solution: false -->
```python
# Test Case 3
assert vowel2num('apple pie') == ('1ppl2 p32', 4)
```

<!-- grade_id: cell-490a3540961d9626, points: 1, grade: true, solution: false -->
```python
# Test Case 4
assert vowel2num('banana') == ('b1n1n1', 3)
```

<!-- grade_id: cell-270bdd28dc17c311, points: 1, grade: true, solution: false -->
```python
# Test Case 5
assert vowel2num('beautiful day') == ('b215t3f5l d1y', 6)
```

<!-- grade_id: cell-ba1ad95eca6b5894, points: 1, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-0c7dbdcf484b2206, points: 1, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-6ec9cc8915d11ba7, points: 1, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-fdd0a9db84cb63fd, points: 1, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-6534005562dcc200, points: 1, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #6 **Sentence Capitalizer** (Total: 10 points)</mark>
สร้างฟังก์ชัน `capitalize_sentences(paragraph)` ซึ่งแปลงตัวอักษรตัวแรกของแต่ละประโยคเป็นตัวพิมพ์ใหญ่ ดังเงื่อนไขต่อไปนี้

1. ฟังก์ชัน `capitalize_sentences(paragraph)` รับพารามิเตอร์เป็นสตริงและทำให้ตัวอักษรตัวแรกของแต่ละประโยคเป็นตัวพิมพ์ใหญ่
   - ประโยคจะถูกกำหนดว่าเป็นลำดับของตัวอักษรที่ลงท้ายด้วยจุด (`.`), เครื่องหมายอัศเจรีย์ (`!`), หรือเครื่องหมายคำถาม (`?`), ตามด้วยช่องว่างหรือลงท้ายด้วยสตริง
2. ฟังก์ชันจะต้อง return คำตอบเป็นสตริง:
   - ประโยคจะต้องถูกจัดรูปแบบโดยตัวอักษรตัวแรกจะเป็นตัวพิมพ์ใหญ่ ในขณะที่ส่วนที่เหลือของข้อความจะคงอยู่ในรูปแบบต้นฉบับ

- ตัวอย่าง `capitalize_sentences('hello world. this is a test! is it working? yes, it is.')` จะได้ `'Hello world. This is a test! Is it working? Yes, it is.'`

<!-- grade_id: cell-78a074b6e25f1ade, points: 0, grade: false, solution: true -->
```python
def capitalize_sentences(paragraph):
    # YOUR CODE HERE
    con = ".,?!"
    para = paragraph.split(" ")
    index_st = 0
    res = []
    for i in range(len(para)):
        if con[0] in para[i] or con[1] in para[i] or con[2] in para[i] or con[3] in para[i]:
            para[index_st] = para[index_st].capitalize()
            index_st = i+1
    print(" ".join(para))
    return " ".join(para)



```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**1 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-fc4273c2f447a29f, points: 1, grade: true, solution: false -->
```python
# Test Case 1
assert capitalize_sentences('hello world. this is a test! is it working? yes, it is.') == 'Hello world. This is a test! is it working? yes, it is.'
```

Output:
```
Hello world. This is a test! Is it working? Yes, It is.
AssertionError                            Traceback (most recent call last)
Cell In[232], line 2
      1 # Test Case 1
----> 2 assert capitalize_sentences('hello world. this is a test! is it working? yes, it is.') == 'Hello world. This is a test! is it working? yes, it is.'

AssertionError:
```

<!-- grade_id: cell-cfc2b6b617f37a71, points: 1, grade: true, solution: false -->
```python
# Test Case 2
assert capitalize_sentences('first sentence. second sentence. third sentence.') == 'First sentence. Second sentence. Third sentence.'
```

Output:
```
First sentence. Second sentence. Third sentence.
```

<!-- grade_id: cell-dd832f379273654d, points: 1, grade: true, solution: false -->
```python
# Test Case 3
assert capitalize_sentences('capitalize me. another test here.') == 'Capitalize me. Another test here.'
```

Output:
```
Capitalize me. Another test here.
```

<!-- grade_id: cell-c5a5ea6742de8bf3, points: 1, grade: true, solution: false -->
```python
# Test Case 4
assert capitalize_sentences('this should work. it really should!') == 'This should work. It really should!'
```

Output:
```
This should work. It really should!
```

<!-- grade_id: cell-2ee7faf1db067a2d, points: 1, grade: true, solution: false -->
```python
# Test Case 5
assert capitalize_sentences('a single sentence.') == 'A single sentence.'
```

Output:
```
A single sentence.
```

<!-- grade_id: cell-1bbb64e2c5d6c546, points: 1, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-9299b959179e1636, points: 1, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-1ac88e5bad1cfede, points: 1, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-c104aae704f6d596, points: 1, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-190a244a8a27da42, points: 1, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #7 **Palindrome Checker** (Total: 10 points)</mark>
สร้างฟังก์ชันตรวจคุณสมบัติ Palindrome ซึ่งมีคำจำกัดความ คือ คำ, วลี, หรือ ลำดับที่อ่านเหมือนกันทั้งจากด้านหน้าและด้านหลัง ดังรายละเอียดต่อไปนี้

1. สร้างฟังก์ชัน `is_palindrome(word)` ที่ตรวจสอบว่าคำที่กำหนดเป็น palindrome หรือไม่
   - ฟังก์ชันพิจารณาตัวพิมพ์ใหญ่พิมพ์เล็กเหมือนกัน และไม่สนใจช่องว่าง เช่น "Racecar" และ "race car" ให้พิจารณาว่าเหมือนกัน
2. ฟังก์ชัน return กลับเป็นค่า boolean
   - `True` หากคำเป็น palindrome
   - `False` หากไม่ใช่

- ตัวอย่าง `is_palindrome('racecar')` จะได้ `True`

<!-- grade_id: cell-86a06d33e59def9c, points: 0, grade: false, solution: true -->
```python
def is_palindrome(word):
    # YOUR CODE HERE
    word = word.lower()
    word =word.replace(" ","")
    word =word.replace("?","")
    return word == word[::-1]
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**1 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-adeadef2394de477, points: 1, grade: true, solution: false -->
```python
# Test Case 1
assert is_palindrome('A man a plan a canal Panama') == True
```

<!-- grade_id: cell-1c41571ef5f7f78d, points: 1, grade: true, solution: false -->
```python
# Test Case 2
assert is_palindrome('racecar') == True
```

<!-- grade_id: cell-8df9302820d4552b, points: 1, grade: true, solution: false -->
```python
# Test Case 3
assert is_palindrome('Hello') == False
```

<!-- grade_id: cell-4e0e48afd72f358c, points: 1, grade: true, solution: false -->
```python
# Test Case 4
assert is_palindrome('Was it a car or a cat I saw?') == True
```

<!-- grade_id: cell-4d3b3749a0ea2793, points: 1, grade: true, solution: false -->
```python
# Test Case 5
assert is_palindrome('No lemon, no melon') == True
```

<!-- grade_id: cell-fc30a6cd95304e3c, points: 1, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-50fb43a46eb39439, points: 1, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-be3e06367a9a6849, points: 1, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-b6e0a0eba379123d, points: 1, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-6b2b5bea86eef2cc, points: 1, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #8 **Anagram Checker** (Total: 10 points)</mark>
สร้างฟังก์ชัน `are_anagrams(word1, word2)`ที่ตรวจสอบว่าคำสองคำที่กำหนดเป็น anagrams ของกันและกันหรือไม่ (anagram คือ คำหรือวลีที่สร้างจากการจัดเรียงตัวอักษรของคำหรือวลีที่แตกต่างกัน)

เงื่อนไขการทำงาน ดังนี้
1. ฟังก์ชันพิจารณาเฉพาะตัวอักษร (เช่น ไม่พิจารณาช่องว่างและเครื่องหมายวรรคตอน) และตัวพิมพ์ใหญ่พิมพ์เล็กให้ผลเหมือนกัน
2. ฟังก์ชัน return เป็นค่า boolean
   - `True` หากคำเป็น anagrams
   - `False` หากไม่ใช่

ตัวอย่าง:
- `are_anagrams('Listen', 'Silent')` จะได้ `True`
- `are_anagrams('Dormitory', 'Dirty Room')` จะได้ `True` เพราะทั้งสองวลีประกอบด้วยตัวอักษรเดียวกัน: `d, i, m, o, o, r, r, t, y`

<!-- grade_id: cell-4c86b3bcb548dd81, points: 0, grade: false, solution: true -->
```python
def are_anagrams(word1, word2):
    # YOUR CODE HERE
    word1 = word1.lower()
    word1 = word1.replace(" ","")
    word2 = word2.lower()
    word2 = word2.replace(" ","")
    return ''.join(sorted(word1)) == ''.join(sorted(word2))
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**1 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-49648dd9501d4d65, points: 1, grade: true, solution: false -->
```python
# Test Case 1
assert are_anagrams('Listen', 'Silent') == True
```

<!-- grade_id: cell-50f9d1dc867a9325, points: 1, grade: true, solution: false -->
```python
# Test Case 2
assert are_anagrams('Triangle', 'Integral') == True
```

<!-- grade_id: cell-6a3db42d9053e8c7, points: 1, grade: true, solution: false -->
```python
# Test Case 3
assert are_anagrams('Apple', 'Papel') == True
```

<!-- grade_id: cell-0bf6f29eca93038c, points: 1, grade: true, solution: false -->
```python
# Test Case 4
assert are_anagrams('Dormitory', 'Dirty Room') == True
```

<!-- grade_id: cell-e9ba7cbe53ed5ba4, points: 1, grade: true, solution: false -->
```python
# Test Case 5
assert are_anagrams('Hello', 'World') == False
```

<!-- grade_id: cell-558809ec86ca5061, points: 1, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-cce3e5e8834ce96f, points: 1, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-f791392fac322ce5, points: 1, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-63b8f31ae76da333, points: 1, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-6a15ae0bf9d37d37, points: 1, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #9 **Caesar Cipher** (Total: 15 points)</mark>
- Caesar cipher เป็นเทคนิคการเข้ารหัสที่ง่าย ซึ่งจะทำการเลื่อนแต่ละตัวอักษรในคำไปยังตำแหน่งที่กำหนดในตัวอักษรของอักษรในภาษาอังกฤษ
- ตัวอักษรแต่ละตัวสามารถเลื่อนได้ทั้ง **left** หรือ **right** เช่น:
  - เมื่อเลื่อน **right** 3 ตำแหน่ง ตัวอักษร `A` จะกลายเป็น `D` และ `Z` จะวนกลับไปเป็น `C`
  - เมื่อเลื่อน **left** 2 ตำแหน่ง ตัวอักษร `C` จะกลายเป็น `A` และ `A` จะวนกลับไปเป็น `Y`
- สร้างฟังก์ชัน `caesar_cipher(word, direction, n)` ที่มีเงื่อนไข ดังนี้
  - `word`: คำที่ต้องการเข้ารหัส ซึ่งประกอบด้วย **ตัวอักษรพิมพ์ใหญ่** 
  - `direction`: ทิศทางการเลื่อน ซึ่งต้องกำหนดว่าเป็น **"left"** หรือ **"right"**
  - `n`: จำนวนตำแหน่งที่ต้องเลื่อนตัวอักษร
- ฟังก์ชัน return ผลลัพธ์เป็นคำที่เข้ารหัส
- ตัวอย่าง
  - `caesar_cipher('HELLO', 'right', 3)` จะได้ผลลัพธ์เป็น `'KHOOR'`
  - `caesar_cipher('WORLD', 'left', 2)` จะได้ผลลัพธ์เป็น `'UPSKB'`

<!-- grade_id: cell-4a8b725e86c3ac80, points: 0, grade: false, solution: true -->
```python
def caesar_cipher(word, direction, n):
    # YOUR CODE HERE
    res = ""
    upcase = "ABCDEFGHIJKLMNOPQRSTUVWXYZ"
    for i in word:
        if direction == "right":
            caesar = (upcase.find(i)+n)%26
            res += upcase[caesar]
        elif direction == "left":
            caesar = (upcase.find(i)-n)%26
            res += upcase[caesar]
    print(res)
    return res

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**1.5 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-1407f8fb7b079e11, points: 1.5, grade: true, solution: false -->
```python
# Test Case 1
assert(caesar_cipher('AABBCC', 'right', 3)=='DDEEFF')
```

Output:
```
DDEEFF
```

<!-- grade_id: cell-120266bf24114be2, points: 1.5, grade: true, solution: false -->
```python
# Test Case 2
assert(caesar_cipher('ABCXYZ', 'right', 15)=='PQRMNO')
```

Output:
```
PQRMNO
```

<!-- grade_id: cell-8cf5aee06f1845f3, points: 1.5, grade: true, solution: false -->
```python
# Test Case 3
assert(caesar_cipher('XYZABC', 'left', 1)=='WXYZAB')
```

Output:
```
WXYZAB
```

<!-- grade_id: cell-431a77d290c7a9c1, points: 1.5, grade: true, solution: false -->
```python
# Test Case 4
assert caesar_cipher('ABC', 'left', 3) == 'XYZ'
```

Output:
```
XYZ
```

<!-- grade_id: cell-8d13f87dbd190766, points: 1.5, grade: true, solution: false -->
```python
# Test Case 5:
assert caesar_cipher('HELLO', 'right', 0) == 'HELLO'
```

Output:
```
HELLO
```

<!-- grade_id: cell-b16732da9223b88e, points: 1.5, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-2d24a60a76ccaa1a, points: 1.5, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-d71883c6ff153321, points: 1.5, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-3e753fb5052c37f2, points: 1.5, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-06dc46e5845741e6, points: 1.5, grade: true, solution: false -->
```python
# Test Case 10
```

---
## <mark>Exercise #10 **Verify Password** (Total: 25 points)</mark>
การสร้าง Password ที่ดี นิยมทำตามกฎ 5 ข้อเพื่อให้มั่นใจว่ามีความปลอดภัย
1. **ความยาว**: รหัสผ่านต้องมีความยาว **อย่างน้อย 8 ตัวอักษร** 
2. **ตัวอักษรพิมพ์เล็ก**: รหัสผ่านต้องมี **อย่างน้อยหนึ่งตัวอักษรพิมพ์เล็ก** (`a-z`)
3. **ตัวอักษรพิมพ์ใหญ่**: รหัสผ่านต้องมี **อย่างน้อยหนึ่งตัวอักษรพิมพ์ใหญ่** (`A-Z`)
4. **ตัวเลข**: รหัสผ่านต้องมี **อย่างน้อยหนึ่งตัวเลข** (`0-9`)
5. **อักขระพิเศษ**: รหัสผ่านต้องมี **อย่างน้อยหนึ่งอักขระพิเศษ** ต่อไปนี้นี้ `-`, `_`, `@`, หรือ `$`

ให้สร้างฟังก์ชัน `verify_pw(password)` ที่ตรวจสอบว่ารหัสผ่านนั้น **ถูกต้อง** หรือ **ไม่ถูกต้อง** ตามกฎทั้ง 5 ข้อนี้

- หากรหัสผ่าน **ถูกต้อง** (ปฏิบัติตามกฎทั้งหมด 5 ข้อ) ให้ส่งคืน `[0]` เพื่อแสดงถึง **รหัสผ่านที่ถูกต้อง**
- หากรหัสผ่าน **ไม่ถูกต้อง** ให้ส่งคืน **รายการปัญหา** ซึ่งแต่ละปัญหาจะแสดงด้วย **หมายเลข** ที่แสดงถึงการไม่ผ่านกฎข้อที่กำหนด ตามลำดับ
  - `1`: รหัสผ่าน **สั้นเกินไป** (น้อยกว่า 8 ตัวอักษร)
  - `2`: รหัสผ่านขาด **ตัวอักษรพิมพ์เล็ก** (`a-z`)
  - `3`: รหัสผ่านขาด **ตัวอักษรพิมพ์ใหญ่** (`A-Z`)
  - `4`: รหัสผ่านขาด **ตัวเลข** (`0-9`)
  - `5`: รหัสผ่านขาด **อักขระพิเศษ** (`-`, `_`, `@`, หรือ `$`)

- ตัวอย่าง
  - `verify_pw('Passw@rd1234')` จะได้ `[0]` แสดงว่ารหัสผ่านถูกต้อง
  - `verify_pw('pw1234')` จะได้ `[1, 3, 5]` แสดงว่ารหัสผ่านสั้นเกินไป ขาดตัวอักษรพิมพ์ใหญ่ และขาดอักขระพิเศษ

- หมายเหตุ
  - ฟังก์ชันต้องประเมินรหัสผ่านและส่งคืนข้อผิดพลาดที่เกี่ยวข้องทั้งหมดหากรหัสผ่านไม่ถูกต้อง
  - สาเหตุของปัญหาจะถูก return ออกมา**ตามลำดับตัวเลข**เสมอ โดยไม่คำนึงถึงลำดับที่พบปัญหา เช่น หากเจอปัญหากฏข้อที่ `3` ก่อน และเจอข้อที่ `1` แต่คำตอบจะได้เป็น `[1, 3]` ตามลำดับตัวเลข

<!-- grade_id: cell-bfaeda995c803778, points: 0, grade: false, solution: true -->
```python
def verify_pw(pw):
    # YOUR CODE HERE
    lower = "abcdefghijklmnopqrstuvwxyz"
    upper = "ABCDEFGHIJKLMNOPQRSTUVWXYZ"
    num = "0123456789"
    sp = "-_@$"
    error = [1,2,3,4,5]
    pw_split = pw.split()
    if len(pw) >= 8:
        error.remove(1)
    for i in pw:
        if i in lower and 2 in error:
            error.remove(2)
        if i in upper and 3 in error:
            error.remove(3)
        if i in num and 4 in error:
            error.remove(4)
        if i in sp and 5 in error:
            error.remove(5)
    print(error)
    return error if len(error) > 0 else [0]



```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**2.5 point per test**</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-a305172dc4f4b573, points: 2.5, grade: true, solution: false -->
```python
# Test Case 1
assert verify_pw('Passw@rd1234') == [0]
```

Output:
```
[]
```

<!-- grade_id: cell-afcd15a3cea10718, points: 2.5, grade: true, solution: false -->
```python
# Test Case 2
assert verify_pw('pw1234') == [1, 3, 5]
```

Output:
```
[1, 3, 5]
```

<!-- grade_id: cell-74e56cc4b38cff43, points: 2.5, grade: true, solution: false -->
```python
# Test Case 3
assert verify_pw('PASSWORD123') == [2, 5]
```

Output:
```
[2, 5]
```

<!-- grade_id: cell-f0488172cce1c94d, points: 2.5, grade: true, solution: false -->
```python
# Test Case 4
assert verify_pw('Password') == [4, 5]
```

Output:
```
[4, 5]
```

<!-- grade_id: cell-048ce53090d6f197, points: 2.5, grade: true, solution: false -->
```python
# Test Case 5
assert verify_pw('Passw0rd!') == [5]
```

Output:
```
[5]
```

<!-- grade_id: cell-77cee0b0a791e99e, points: 2.5, grade: true, solution: false -->
```python
# Test Case 6
```

<!-- grade_id: cell-1812475114e6be0b, points: 2.5, grade: true, solution: false -->
```python
# Test Case 7
```

<!-- grade_id: cell-2ed2efc2964c50c4, points: 2.5, grade: true, solution: false -->
```python
# Test Case 8
```

<!-- grade_id: cell-d6093833f33403f5, points: 2.5, grade: true, solution: false -->
```python
# Test Case 9
```

<!-- grade_id: cell-18da3f03c008194f, points: 2.5, grade: true, solution: false -->
```python
# Test Case 10
```