Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# SECTION 2 ONLY
Name = "Thanakrit Samoena"

Student ID = "67070060"

# WEEK 10: Lists and Tuples - Exercise

<mark>**DO NOT IMPORT EXTERNAL LIBRARY**</mark>

---
## <mark>Exercise #1 (Total: 5 points) </mark>

Complete the `sum_and_average` function takes a list of numbers as input and calculates two values:

- The total sum of the numbers in the list.
- The average of the numbers in the list.

**Parameters:**
- `list1`: A list of numerical values (integers or floats).

**Returns:** A tuple containing two values:
- The sum of all the numbers in the list.
- The average value of the numbers in the list.

**Example:**
- `list1 = [1, 2, 3, 4, 5]` returns a sum of `15` and an average of `3.0`.

<!-- grade_id: cell-df0ae4a6d6764a38, points: 0, grade: false, solution: true -->
```python
def sum_and_average(numbers):

    # YOUR CODE HERE
    return sum(numbers),sum(numbers)/len(numbers)
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-212615214f873668, points: 1, grade: true, solution: false -->
```python
# Test 1: 1 point
list1 = [-1, 0, 1, 2, 3]
assert(sum_and_average(list1) == (5, 1.0))
```

<!-- grade_id: cell-768672447efbf053, points: 1, grade: true, solution: false -->
```python
# Test 2: 1 point (Partially hidden)
list1 = [0, 0, 0, 0]

```

<!-- grade_id: cell-79500d2f0eb0d298, points: 1, grade: true, solution: false -->
```python
# Test 3: 1 point (Partially hidden)
list1 = [7]

```

<!-- grade_id: cell-5b0baef57d9243f6, points: 1, grade: true, solution: false -->
```python
# Test 4: 1 point (Completely hidden)
```

<!-- grade_id: cell-be41a1fc3a531cc9, points: 1, grade: true, solution: false -->
```python
# Test 5: 1 point (Completely hidden)
```

## <mark>Exercise #2 (15 points) </mark>

Complete the `missing_number` function to:
1. Find the missing number in a list containing numbers from `1` to `N`.
2. Return a list that includes the missing number in sequence.

**Example:**

- `list2 = [1, 4, 5, 6, 3]`
- `missed_number, new_list = missing_number(list2)`

**Expected Results:**

- The missing number is `missed_number`, which should be `2`.
- The complete list is `new_list`, which should be `[1, 2, 3, 4, 5, 6]`.

<!-- grade_id: list_missing_number1, points: 0, grade: false, solution: true -->
```python
def missing_number(items):

    # YOUR CODE HERE
    items.sort()
    items.insert(0,0)
    for i in range(len(items)):
        if items[i] != i:
            missed_number = i
            items.insert(i,i)
            break
    return missed_number, items[1:] # return one missing number and a completed list.
list2 = [4, 1, 3, 5]
missing_number(list2)
```

Output:
```
(2, [1, 2, 3, 4, 5])
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-4da84509ff99a1cd, points: 3, grade: true, solution: false -->
```python
# Test 1: 3 points
list2 = [3, 1, 4, 5]
assert(missing_number(list2)) == (2, [1, 2, 3, 4, 5])
```

<!-- grade_id: cell-7bcda9d59449c806, points: 3, grade: true, solution: false -->
```python
# Test 2: 3 points (Partially hidden)
list2 = [5, 1, 2, 4]

```

<!-- grade_id: cell-06baf4dad095a3f4, points: 3, grade: true, solution: false -->
```python
# Test 3: 3 points (Partially hidden)
list2 = []

```

<!-- grade_id: list_missing_number_test2, points: 3, grade: true, solution: false -->
```python
# Test 4: 3 points (Completely hidden)

```

<!-- grade_id: list_missing_number_test1, points: 3, grade: true, solution: false -->
```python
# Test 5: 3 points (Completely hidden)

```

---
## <mark>Exercise #3 (Total: 20 points) </mark>
Implement a `sort_unique_tuple` function that:

- Takes a tuple as input.
- Sorts its elements.
- Returns a new tuple containing only the unique elements.

**Parameters:**
- `items`: A tuple of integers.

**Returns:**
- A tuple with elements sorted in ascending order and duplicates removed.

**Example:**
- `tuple3 = (4, 1, 3, 3, 4)`
- `sorted_unique_tuple = sort_unique_tuple(tuple3)`

The expected output for `sorted_unique_tuple` is `(1, 3, 4)`.

<!-- grade_id: tuple_sort1, points: 0, grade: false, solution: true -->
```python
def sort_unique_tuple(items):
    # YOUR CODE HERE
    li = list(items)
    for i in li:
        if li.count(i) > 1:
            li.remove(i)
    return tuple(li)
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-97505c998d310ee0, points: 4, grade: true, solution: false -->
```python
# TEST 1: 4 points
tuple3 = (4, 1, 3, 3, 4)
assert(sort_unique_tuple(tuple3) == (1, 3, 4))
```

<!-- grade_id: cell-d564b3befcdb895a, points: 4, grade: true, solution: false -->
```python
# TEST 2: 4 points (Partially hidden)
tuple3 = (10, 5, 7, 2, 8)

```

<!-- grade_id: cell-e2c72be9c1dfc2d9, points: 4, grade: true, solution: false -->
```python
# TEST 3: 4 points (Partially hidden)
tuple3 = (-1, -3, 2, 1, -1, 3, 2)

```

<!-- grade_id: tuple_sort_test2, points: 4, grade: true, solution: false -->
```python
# TEST 4: 4 points (Completely hidden)

```

<!-- grade_id: tuple_sort_test1, points: 4, grade: true, solution: false -->
```python
# TEST 5: 4 points (Completely hidden)

```

---
## <mark>Exercise # 4 (Total: 20 points) </mark>

Complete a function `num_frequency` that takes a list of numbers and returns a list with two sub-lists:

- The first sub-list contains unique numbers from the input list, sorted in ascending order (low to high).
- The second sub-list contains the counts of each unique number, matching the order in the first sub-list.

**Parameters:**

- `items`: A list of numerical values (integers or floats).

**Returns:** A list with two sub-lists:
- The first sub-list: Unique numbers in ascending order.
- The second sub-list: The frequency of each unique number.

**Example:**
- `list4 = [2, 1, 1, 2, 2, 3, 3, 3, 4, 3]`
- `list_freq = num_frequency(list4)`

**Expected Output:**

`[[1, 2, 3, 4], [2, 3, 4, 1]]`

- The first sub-list `[1, 2, 3, 4]` shows unique numbers in order.
- The second sub-list `[2, 3, 4, 1]` shows the counts of each unique number.


<!-- grade_id: list_frequency1, points: 0, grade: false, solution: true -->
```python
def num_frequency(items):
    # YOUR CODE HERE
    count = []
    newcount = []
    unili = []
    items.sort()
    for i in items:
        if items.count(i) > 1:
            if i not in unili:
                unili.append(i)
        else:
            unili.append(i)
    for i in items:
        count.append(items.count(i))
    for i in count:
        if count.count(i) > 1:
            if i not in newcount:
                newcount.append(i)
        else:
            newcount.append(i)
    print(unili,newcount)
    
    return [unili,newcount]

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-8c366abb439411e0, points: 4, grade: true, solution: false -->
```python
# TEST 1: 4 points
list4 = [2, 1, 1, 2, 2, 3, 3, 3, 4, 3]
assert(num_frequency(list4) == [[1, 2, 3, 4], [2, 3, 4, 1]])
```

Output:
```
[1, 2, 3, 4] [2, 3, 4, 1]
```

<!-- grade_id: cell-bb52ad682cd78169, points: 4, grade: true, solution: false -->
```python
# TEST 2: 4 points (Partially hidden)
list4 = [5, 3, 1, 4, 2]

```

<!-- grade_id: cell-3c3d6be5bde5a1d9, points: 4, grade: true, solution: false -->
```python
# TEST 3: 4 points (Partially hidden)
list4 = [7, 7, 7, 8, 8, 9]

```

<!-- grade_id: list_frequency_test2, points: 4, grade: true, solution: false -->
```python
# TEST 4: 4 points (Completely hidden)

```

<!-- grade_id: list_frequency_test1, points: 4, grade: true, solution: false -->
```python
# TEST 5: 4 points (Completely hidden)

```

---
## <mark>Exercise #5 (Total: 20 points) </mark>
Complete a function `split_numbers` that takes a `list` or `tuple` of numbers and separates them into two collections:

- **Even numbers**: Collect these in a `list` (unsorted).
- **Odd numbers**: Collect these in a `tuple` (unsorted).

**Parameters:**
- numbers: A `list` or `tuple` of integers.

**Returns:** A tuple containing:
- A `list` of even numbers.
- A `tuple` of odd numbers.

**Example:**
- `list5 = [2, 3, 4, 3, 2, 2, 3, 1, 2, 3]`
- `even_numbers, odd_numbers = split_numbers(list5)`

**Output:**
- `even_numbers: [2, 4, 2, 2, 2]`
- `odd_numbers: (3, 3, 3, 1, 3)`
- The `even_numbers` list contains all even numbers from the input.
- The `odd_numbers` tuple contains all odd numbers from the input.

<!-- grade_id: split_ev_odd1, points: 0, grade: false, solution: true -->
```python
def split_numbers(items):
    # YOUR CODE HERE
    odd = []
    even = []
    for j in items:
        if j % 2 == 0:
            even.append(j)
        else:
            odd.append(j)
    return (even,tuple(odd))
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-17d481df73c6bfff, points: 4, grade: true, solution: false -->
```python
# TEST 1: 4 points
list5 = [2, 3, 4, 3, 2, 2, 3, 1, 2, 3]
assert(split_numbers(list5) == ([2, 4, 2, 2, 2], (3, 3, 3, 1, 3)))
```

<!-- grade_id: cell-9dd9f283f5e7e48c, points: 4, grade: true, solution: false -->
```python
# TEST 2: 4 points (Partially hidden)
list5 = [6, 8, 10, 2, 4]

```

<!-- grade_id: cell-60eb922b330dcc8f, points: 4, grade: true, solution: false -->
```python
# TEST 3: 4 points (Partially hidden)
list5 = [10, 21, 32, 43, 54, 65]

```

<!-- grade_id: split_ev_odd_test2, points: 4, grade: true, solution: false -->
```python
# TEST 4: 4 points (Completely hidden)

```

<!-- grade_id: split_ev_odd_test1, points: 4, grade: true, solution: false -->
```python
# TEST 5: 4 points (Completely hidden)

```

---
## <mark>Exercise #6 (Total: 20 points) </mark>
Complete a `sales_report` function that processes a dataset of car sales information from various countries in 2021 and provides a summary for a specified country.

The function should:

**Input:** 
- A dataset `car_sales` consisting of tuples where each tuple represents:
    - The country of sale (string)
    - The car brand (string)
    - The number of units sold (integer)
- A `country` (string) specifying the country for which the report is required.

**Output:**
- A tuple with two values:
    - `brand`: A list of unique car brands sold in the specified country.
    - `total_sales`: The total number of units sold for all car brands in the specified country.

**Example:**
```python
car_sales = [('Japan', 'Toyota', 1933099), 
             ('US', 'Ford', 1804824), 
             ('US', 'Chevrolet', 1422031), 
             ('Japan', 'Honda', 1308476), 
             ('Japan', 'Nissan', 919090), 
             ('US', 'Jeep', 768713), 
             ('South Korea', 'Hyundai', 726715), 
             ('South Korea', 'Kia', 677494), 
             ('US', 'Ram', 647329), 
             ('Japan', 'Subaru', 598480)]
```

`brand, total_sales = sales_report(car_sales, 'US')`

**Expected Output:**

- `brand` should be a list of car brands sold in the `'US'`: `['Ford', 'Chevrolet', 'Jeep', 'Ram']`.
- `total_sales` should be the total number of units sold in the `'US'`: `4642897`.

*Note:* [Full-Year 2021 Auto Sales Report](https://www.carpro.com/blog/full-year-2021-sales-report-with-most-brands-reporting)

<!-- grade_id: sale_report1, points: 0, grade: false, solution: true -->
```python
def sales_report(items, sales_by_country): 
    # YOUR CODE HERE
    res = 0
    car = []
    for i in items:
        if sales_by_country in i:
            car.append(i[1])
            res += i[2]
    return (car,res)
```

```python
car_sales = [('Japan', 'Toyota', 1933099), 
             ('US', 'Ford', 1804824), 
             ('US', 'Chevrolet', 1422031), 
             ('Japan', 'Honda', 1308476), 
             ('Japan', 'Nissan', 919090), 
             ('US', 'Jeep', 768713), 
             ('South Korea', 'Hyundai', 726715), 
             ('South Korea', 'Kia', 677494), 
             ('US', 'Ram', 647329), 
             ('Japan', 'Subaru', 598480)]
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-1ef808f9982a77c2, points: 5, grade: true, solution: false -->
```python
# TEST 1: 5 points
assert(sales_report(car_sales, 'US') == (['Ford', 'Chevrolet', 'Jeep', 'Ram'], 4642897))

```

<!-- grade_id: cell-274964253149b99c, points: 5, grade: true, solution: false -->
```python
# TEST 2: 5 points
country = "Canada"
assert(sales_report(car_sales, country) == ([], 0))

```

<!-- grade_id: sale_report_test2, points: 5, grade: true, solution: false -->
```python
# TEST 3: 5 points (Completely hidden)

```

<!-- grade_id: sale_report_test1, points: 5, grade: true, solution: false -->
```python
# # TEST 4: 5 points (Completely hidden)

```