Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# WEEK 06: Function - Exercise

**Name**: Thanakrit Samoena

**ID**: 67070060

**Program**: AIT

**Write you codes after `# YOUR CODE HERE`**

**BUILT-IN FUNCTIONS (e.g. `min` function)or EXTERNAL LIBRARIES ARE NOT ALLOWED.**

---
<span style='background :yellow' > Execise #1 </span>.

Define a function to calculate the area of any circle where the area of circle is given by 

$ A = \pi r^2$,

where $\pi = 3.14$ and $r$ is the circle radius and $r > 0$.

Note. If $r \le 0$, the result must be `-321`

<!-- grade_id: circ_area_fnc, points: 0, grade: false, solution: true -->
```python
Pi = 3.14
def circ_area(Pi, radius):
    """ Return the area of circle."""
    # YOUR CODE HERE
    if radius <= 0:
        return -321
    else :
        return Pi*radius**2

```

<span style='background :cyan'>
DO NOT CHANGE THE FOLLOWING LINES
</span>

Automatic Grading with the following tests. The correct results will get total 10 points.

<!-- grade_id: circ_area_fnc_test1, points: 5, grade: true, solution: false -->
```python
# Test-1: 5 Points
assert circ_area(Pi, 5) == 78.5
assert circ_area(Pi, 2) == 12.56
```

<!-- grade_id: circ_area_fnc_test2, points: 5, grade: true, solution: false -->
```python
# Test-2: 5 Points
assert circ_area(Pi, -1) == -321   # false condition
assert circ_area(Pi, 0) == -321    # false condition
```

---
<span style='background :yellow' > Execise #2 </span>.

Define a function to calculate the perimeter of a triangle, where the perimeter of the triangle is given by 

$P = a + b + c$, where $a, b,$ and $c$ are the sides of the triangle.

Note. If $a, b,$ or $c \le 0$, the result must be `-321`

<!-- grade_id: cell-5c1f97048a90d6de, points: 0, grade: false, solution: true -->
```python
def rect_triangle(a, b, c):
    """Return the perimeter of a triangle."""

    # YOUR CODE HERE
    if a <= 0 or b <= 0 or c <= 0:
        return -321
    else:
        return a+b+c
```

<span style='background :cyan'>
DO NOT CHANGE THE FOLLOWING LINES
</span>

Automatic Grading with the following tests. The correct results will get total 10 points.

<!-- grade_id: cell-5aee4fbb54e50261, points: 5, grade: true, solution: false -->
```python
# Test-1: 5 Points
assert rect_triangle(1, 2, 3) == 6
assert rect_triangle(1, -1, 5) == -321
```

<!-- grade_id: cell-0a3639ce67a07d95, points: 5, grade: true, solution: false -->
```python
# Test-2: 5 Points
assert rect_triangle(1, -4, 10) == -321
assert rect_triangle(99, 99, 99) == 297
```

---
<span style='background :yellow' > Execise #3 </span>.

Define a function to calculate the volume of any cube where the volume of cube is given by 

$ V = a^3$,

where $a$ is the cube edge and $a$ >= 0$.

Note. If $a \le 0$, the result must be `-321`

<!-- grade_id: cell-a2cec35c2bfc88bb, points: 0, grade: false, solution: true -->
```python
def cube_volume(a):
    """Return the volume of a cube."""

    # YOUR CODE HERE
    if a <= 0:
        return -321
    else:
        return a**3

```

<span style='background :cyan'>
DO NOT CHANGE THE FOLLOWING LINES
</span>

Automatic Grading with the following tests. The correct results will get total 10 points.

<!-- grade_id: cell-6c9ff2c4bfa08224, points: 5, grade: true, solution: false -->
```python
# Test-1: 5 Points
assert cube_volume(2) == 8
assert cube_volume(-8) == -321
```

<!-- grade_id: cell-fbc875b73fbe2797, points: 5, grade: true, solution: false -->
```python
# Test-2: 5 Points
assert cube_volume(0) == -321
assert cube_volume(4) == 64
```

---
<span style='background :yellow' > Execise #4 </span>.


Create a function that returns the **Minimum** value of three numbers.

- Input: receving three numbers
- Output: returning the minimnum of input numbers

<!-- grade_id: min_fnc, points: 0, grade: false, solution: true -->
```python
def min_value(a, b, c):
    """ Return the minimum of three values."""
    # YOUR CODE HERE
    min = a
    if a < b:
        min = a
    else:
        min = b
        if c < b:
            min = c
    return min
    
```

<span style='background :cyan'>
DO NOT CHANGE THE FOLLOWING LINES
</span>

Automatic Grading with the following tests. The correct results will get total 15 points.

<!-- grade_id: min_fnc_test1, points: 5, grade: true, solution: false -->
```python
# Test-1: 5 points
assert min_value(-2,5,7) == -2
```

<!-- grade_id: min_fnc_test2, points: 5, grade: true, solution: false -->
```python
# Test-2: 5 points
assert min_value(5,-2,7) == -2
```

<!-- grade_id: min_fnc_test3, points: 5, grade: true, solution: false -->
```python
# Test-3: 5 points
assert min_value(7,5,-2) == -2
```

---
<span style='background :yellow' > Execise #5 </span>.

The definition of **Factorial** of a whole number `n` is defined as the product of that number with every whole number less than or equal to `n` till 1, given by

$n! = n * (n-1)!$.

For example, $4! = 4 * 3 * 2 * 1 = 24$. Note that $0! = 1$ and $1! = 1$.

Create a function to calculate a factorial of any integer number `n` where `n` $\ge 0$.

Note. If $n < 0$, the result must be `-123`

<!-- grade_id: fac_fnc, points: 0, grade: false, solution: true -->
```python
def fac(n):
    """ Return the factorial of integer number"""
    # YOUR CODE HERE
    result = 1
    if n < 0:
        return -123
    elif not n:
        return 1
    else:
        for i in range(1,n+1):
            result *= i
        return result
fac(4)

```

Output:
```
24
```

<span style='background :cyan'>
DO NOT CHANGE THE FOLLOWING LINES
</span>

Automatic Grading with the following tests. The correct results will get total 25 points.

<!-- grade_id: fac_fnc_test1, points: 5, grade: true, solution: false -->
```python
# # Test-1: 5 points
assert fac(0) == 1
assert fac(1) == 1
```

<!-- grade_id: fac_fnc_test2, points: 15, grade: true, solution: false -->
```python
# Test-2: 15 points
assert fac(22) == 1124000727777607680000
```

<!-- grade_id: fac_fnc_test3, points: 5, grade: true, solution: false -->
```python
# Test-3: 5 points
assert fac(-1) == -123
```

---
<span style='background :yellow' > Execise #6 </span>.

**Restaurant Total Cost**

Create functions that calculate the total cost of eating at a restaurant, where the total cost can be computed as follows:

`food_cost`: the cost of the food you eat.

`tax_rate`: the tax based on the food cost.

`service_charge_percentage`: the service charge imposed by the restaurant based on the food cost.

`is_take_home`: if it is `True`, the `service_charge_percentage` will be excluded.

The `total_cost` is calculated as the sum of `food_cost`, `tax_rate`, and `service_charge_percentage`.

<!-- grade_id: cell-a4452537dc8d3645, points: 0, grade: false, solution: true -->
```python
def calculate_total_cost(food_cost, tax_rate, service_charge_percentage, is_take_home):
    """ Return the total cost of eating at restaurant"""
    # YOUR CODE HERE
    if not is_take_home:
        total_cost = food_cost+(food_cost*tax_rate)+(food_cost*service_charge_percentage)
        return  total_cost
    else :
        total_cost = food_cost+(food_cost*tax_rate)
        return  total_cost
    
```

<span style='background :cyan'>
DO NOT CHANGE THE FOLLOWING LINES
</span>

Automatic Grading with the following tests. The correct results will get total 30 points.

<!-- grade_id: cell-2a7759de82f4f6df, points: 7.5, grade: true, solution: false -->
```python
# Test4-1: 7.5 points
food_cost = 100
tax_rate = 0.07  # 7% tax rate
service_charge_percentage = 0.10  # 10% service charge
is_take_home = False    # eat at resterant

# Calculate and display the total cost
total_cost = calculate_total_cost(food_cost, tax_rate, service_charge_percentage, is_take_home)

assert total_cost == 117.0
```

<!-- grade_id: cell-fed1dfc444cfe0ff, points: 7.5, grade: true, solution: false -->
```python
# Test4-2: 7.5 points
food_cost = 100
tax_rate = 0.07  # 7% tax rate
service_charge_percentage = 0.10  # 10% service charge
is_take_home = True    # eat at home

# Calculate and display the total cost
total_cost = calculate_total_cost(food_cost, tax_rate, service_charge_percentage, is_take_home)

assert total_cost == 107.0
```

<!-- grade_id: cell-f546fe1c2a30ceb5, points: 7.5, grade: true, solution: false -->
```python
# Test4-3: 7.5 points
food_cost = 5
tax_rate = 0.07  # 7% tax rate
service_charge_percentage = 0.10  # 10% service charge
is_take_home = False    # eat at resterant

# Calculate and display the total cost
total_cost = calculate_total_cost(food_cost, tax_rate, service_charge_percentage, is_take_home)
assert total_cost == 5.85
```

<!-- grade_id: cell-d04d0387110ddbf9, points: 7.5, grade: true, solution: false -->
```python
# Test4-4: 7.5 points
food_cost = 5
tax_rate = 0.07  # 7% tax rate
service_charge_percentage = 0.10  # 10% service charge
is_take_home = True    # eat at home

# Calculate and display the total cost
total_cost = calculate_total_cost(food_cost, tax_rate, service_charge_percentage, is_take_home)
assert total_cost == 5.35
```