Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# SECTION 2 ONLY
Name = "Thanakrit Samoena"

Student ID = "67070060"

# WEEK 11: Dictionary and Set - Exercise

<mark>**EXTERNAL LIBRARIES ARE NOT ALLOWED.**</mark>

---

## <mark>Exercise #1 (Total: 5 points) </mark>

สร้าง set ที่ประกอบด้วยชื่อของนักเรียน โดย set อยู่ในรูปแบบดังนี้: `student_set = {student_name}`

รายชื่อของนักเรียน 10 คน คือ

- `Somchai`
- `Somsak`
- `Somying`
- `Sompong`
- `Somsri`
- `Somjit`
- `Somchoke`
- `Somsamorn`
- `Somrak`
- `Somtawil`

<!-- grade_id: cell-579bb060eade19e4, points: 0, grade: false, solution: true -->
```python
student_set = {"Somchai"
,"Somsak"
,"Somying"
,"Sompong"
,"Somsri"
,"Somjit"
,"Somchoke"
,"Somsamorn"
,"Somrak"
,"Somtawil"}    # MODIFY YOUR CODE HERE

# YOUR CODE HERE

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-62eebb257f9dcd36, points: 1, grade: true, solution: false -->
```python
# Test 1: 1 point
assert('Somjit' in student_set)
```

<!-- grade_id: cell-ebec89f8a2cc5d5a, points: 1, grade: true, solution: false -->
```python
# Test 2: 1 point
assert('Somporn' not in student_set)
```

<!-- grade_id: cell-0178c31a43ad5587, points: 1, grade: true, solution: false -->
```python
# Test 3: 1 point
assert('Somsamorn' in student_set)
```

<!-- grade_id: cell-f6bafc4d5fa19393, points: 1, grade: true, solution: false -->
```python
# Test 4: 1 point (Completely hidden)
```

<!-- grade_id: cell-20f58f91a9eb9f61, points: 1, grade: true, solution: false -->
```python
# Test 5: 1 point (Completely hidden)
```

---
## <mark>Exercise #2 (Total: 5 points) </mark>

สร้าง dictionary ที่เก็บชื่อของนักเรียนและคะแนนในรูปแบบดังนี้ `student_grades = {student_name: grade}`

- `Somchai: 85`
- `Somsak: 90`
- `Somying: 78`
- `Sompong: 92`
- `Somsri: 88`
- `Somjit: 77`
- `Somchoke: 68`
- `Somsamorn: 49`
- `Somrak: 70`
- `Somtawil: 62`

<!-- grade_id: cell-07434ca88b2311d4, points: 0, grade: false, solution: true -->
```python
student_grades = {"Somchai": 85
,"Somsak": 90
,"Somying": 78
,"Sompong": 92
,"Somsri": 88
,"Somjit": 77
,"Somchoke": 68
,"Somsamorn": 49
,"Somrak": 70
,"Somtawil": 62}    # MODIFY YOUR CODE HERE

# YOUR CODE HERE

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-eb063c940361988f, points: 1, grade: true, solution: false -->
```python
# Test 1: 1 point
assert(student_grades['Somchai'] == 85)
```

<!-- grade_id: cell-041f120a03ade7cc, points: 1, grade: true, solution: false -->
```python
# Test 2: 1 point
assert(student_grades['Somjit'] == 77)
```

<!-- grade_id: cell-f9c9f118805840be, points: 1, grade: true, solution: false -->
```python
# Test 3: 1 point
assert(student_grades['Somtawil'] == 62)
```

<!-- grade_id: cell-e05f0249e4145c99, points: 1, grade: true, solution: false -->
```python
# Test 4: 1 point (Completely hidden)
```

<!-- grade_id: cell-1073f64daa62ae0f, points: 1, grade: true, solution: false -->
```python
# Test 5: 1 point (Completely hidden)
```

---
## <mark>Exercise #3 (Total: 10 points) </mark>

สร้างฟังก์ชัน `manage_grades(operation, job_name, salary=None)` เพื่อจัดการรายการของ Dictionary ที่แสดงถึงตำแหน่งงานและเงินเดือน

- ฟังก์ชันรับพารามิเตอร์ 3 ตัว ดังนี้
  1. `operation`: string ที่ระบุ operation ที่ต้องทำ คือ `'add'` เพื่อเพิ่มตำแหน่งงาน หรือ `'remove'` เพื่อลบตำแหน่งงานออก
  2. `job_name`: ชื่อตำแหน่งงาน
  3. `salary`: เงินเดือนสำหรับตำแหน่งงาน (ใช้เฉพาะเมื่อเพิ่มตำแหน่งงาน)

- ฟังก์ชันจะแก้ไข Dictionary `job_salary` ในรูปแบบ global และจะคืนค่า `1` หากการดำเนินการสำเร็จ หรือ `-1` หากล้มเหลว

#### ตัวอย่าง:

- `manage_grades('add', 'Pilot', 345000)` จะเพิ่มตำแหน่งงาน `Pilot` พร้อมเงินเดือน `345000` และคืนค่า `1` แสดงถึงความสำเร็จ
- `manage_grades('remove', 'Manager')` จะคืนค่า `-1` เนื่องจากตำแหน่งงาน `Manager` ไม่มีใน Dictionary

<!-- grade_id: cell-73b9de3139a76c44, points: 0, grade: false, solution: true -->
```python
job_salary = {'Engineer': 50000, 'Doctor': 100000, 'Lawyer': 150000, 'Actor': 200000}

def manage_grades(operation, job_name, salary = None):
    global job_salary     # Not necessary, but included for illustrative purposes only
    
    # YOUR CODE HERE
    if operation == "add":
        job_salary[job_name] = salary
        return 1
    elif operation == "remove":
        if job_name in job_salary:
            job_salary.pop(job_name)
            return 1
        else:
            return -1

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: cell-29bc75b3f2f48986, points: 2, grade: true, solution: false -->
```python
# Test 1: 2 points
manage_grades('add', 'CEO', 500000)
assert(manage_grades('add', 'CEO', 500000) == 1)
```

<!-- grade_id: cell-96b72bce6b3b8297, points: 2, grade: true, solution: false -->
```python
# Test 2: 2 points
assert(job_salary['CEO'] == 500000)
```

<!-- grade_id: cell-c52ee723326b869e, points: 2, grade: true, solution: false -->
```python
# Test 3: 2 points
assert(manage_grades('remove', 'Doctor') == 1)
```

<!-- grade_id: cell-7d8168645b397e6f, points: 2, grade: true, solution: false -->
```python
# Test 4: 2 points (Partially  hidden)
manage_grades('remove', 'Doctor')   # Remove 'Doctor' again.

```

Output:
```
-1
```

<!-- grade_id: cell-94cb79e484dc05ef, points: 2, grade: true, solution: false -->
```python
# Test 5: 2 points (Completely hidden)

```

---

## <mark>Exercise #4 (Total: 50 points) </mark>

สร้างฟังก์ชัน **search engine** พื้นฐานที่สามารถแสดงผลลัพธ์ตาม keyword ที่กำหนด
- หากพบคำที่ตรงกับ keyword ให้ return ผลลัพธ์ทั้งหมด
- หากไม่พบคำที่ตรงกับ keyword ให้ return `{'Not Found'}`

**หมายเหตุ:** การค้นหา **ไม่คำนึงถึงตัวพิมพ์ใหญ่และพิมพ์เล็ก** ดังนั้น `ThisKeyword` จะถือเป็นคำเดียวกับ `thiskeyword`

### ตัวอย่าง:

1. `print(my_search_engine(swift_code, 'Caribbean'))`
   - **ผลลัพธ์:** `{'Caribbean Development Bank'}`

2. `print(my_search_engine(swift_code, 'internatIONAL'))`
   - **ผลลัพธ์:** `{'International Fund for Agricultural Development', 'International Development Association', 'International Finance Corporation', 'International Bank for Reconstruction and Development'}`

3. `print(my_search_engine(swift_code, 'abc'))`
   - **ผลลัพธ์:** `{'Not Found'}`   

<!-- grade_id: seach_func, points: 0, grade: false, solution: true -->
```python
# USE THIS DICTIONARY AND SET TO ANSER THIS QUESTIONS

swift_code = {'AFDBKEN1': 'African Development Bank',
              'ASDBPHMM': 'Asian Development Bank',
              'CDEVBBBB': 'Caribbean Development Bank',
              'EBRDGB2L': 'European Bank for Reconstruction and Development',
              'IBRDUS33': 'International Bank for Reconstruction and Development',
              'IFCWUS33': 'International Finance Corporation',
              'NOINFIHH': 'Nordic Investment Bank',
              'IFADITRR': 'International Fund for Agricultural Development',
              'IDAWUS33': 'International Development Association',
              'UNESFRP1': 'United Nations Educational, Scientific and Cultural'}

European_results = {'European Bank for Reconstruction and Development'}
Asia_results = {'Not Found'}
asian_results = {'Asian Development Bank'}
bank_results = {'Nordic Investment Bank', 
                'International Bank for Reconstruction and Development', 
                'African Development Bank', 'Asian Development Bank', 
                'European Bank for Reconstruction and Development', 
                'Caribbean Development Bank'}

def my_search_engine(dictionary, search_keyword):
    # YOUR CODE HERE
    dic_ans = []
    search_keyword = search_keyword.capitalize()
    for i in dictionary.values():
        if search_keyword in i.split():
            dic_ans.append(i)
    if len(dic_ans) < 1:
        dic_ans.append("Not Found")
    print(set(dic_ans))
    return set(dic_ans)
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

<!-- grade_id: search_1, points: 10, grade: true, solution: false -->
```python
# Test 1: 10 points
assert(my_search_engine(swift_code, 'European') == European_results)
```

Output:
```
{'European Bank for Reconstruction and Development'}
```

<!-- grade_id: search_2, points: 10, grade: true, solution: false -->
```python
# Test 2: 10 points
assert(my_search_engine(swift_code, 'Asia') == Asia_results)
```

Output:
```
{'Not Found'}
```

<!-- grade_id: search_3, points: 10, grade: true, solution: false -->
```python
# Test 3: 10 points
assert(my_search_engine(swift_code, 'asian') == asian_results)
```

Output:
```
{'Asian Development Bank'}
```

<!-- grade_id: search_4, points: 20, grade: true, solution: false -->
```python
# Test 4: 20 points (Completely hidden)

```

---

## <mark>Exercise #5 (Total: 30 points) </mark>

สร้างฟังก์ชัน `words_count(words)` เพื่อทำการนับความถี่ของแต่ละคำและเครื่องหมายวรรคตอนใน string จาก `words` ที่กำหนด

**คำแนะนำ:**
1. ฟังก์ชันควรจัดการกับ string ที่ป้อนมาในลักษณะที่**ไม่คำนึงถึงตัวพิมพ์ใหญ่และพิมพ์เล็ก** เช่น `This_Word` ถือเป็นคำเดียวกับ `this_word`
2. ฟังก์ชัน return ผลลัพธ์ในรูปแบบ Dictionary ที่มีรายละเอียด ดังนี้
   - **key** คือ คำและเครื่องหมายวรรคตอนที่ไม่ซ้ำกัน (ในรูปแบบตัวพิมพ์เล็ก)
   - **value** คือ จำนวนครั้งที่แต่ละ **key** ปรากฏใน string

**ตัวอย่าง:**

ให้สตริงที่ป้อนมา: `"It was good , it was bad , it was hot , and it was cold."`

ผลลัพธ์ควรเป็น: `{'it': 4, 'was': 4, 'good': 1, ',': 3, 'bad': 1, 'hot': 1, 'and': 1, 'cold.': 1}`

**หมายเหตุ:**
- คำและเครื่องหมายวรรคตอนจะได้รับการจัดการแยกกัน ดังนั้น นับเครื่องหมายจุลภาค (`,`) และจุด (`.`) แยกกัน
- ตรวจสอบให้แน่ใจว่าเครื่องหมายวรรคตอนที่ติดอยู่กับคำ (เช่น `cold.`)

<!-- grade_id: words_count_func, points: 0, grade: false, solution: true -->
```python
def words_count(words):
    # YOUR CODE HERE
    words = words.lower().split()
    dic_word = {}
    for i in words:
        if i not in dic_word:
            dic_word[i] = words.count(i)
    print(dic_word)
    return dic_word

```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

<mark>**Some tests are hidden**</mark>

```python
# The opening paragraph from "A Tale of Two Cities" by Charles Dickens.
words = 'It was the best of times , it was the worst of times , it was the age of wisdom , it was the age of foolishness , it was the epoch of belief , it was the epoch of incredulity , it was the season of Light , it was the season of Darkness , it was the spring of hope , it was the winter of despair , we had everything before us , we had nothing before us , we were all going direct to Heaven , we were all going direct the other way — in short , the period was so far like the present period , that some of its noisiest authorities insisted on its being received , for good or for evil , in the superlative degree of comparison only.'
word_count_dict = words_count(words)
```

Output:
```
{'it': 10, 'was': 11, 'the': 14, 'best': 1, 'of': 12, 'times': 2, ',': 17, 'worst': 1, 'age': 2, 'wisdom': 1, 'foolishness': 1, 'epoch': 2, 'belief': 1, 'incredulity': 1, 'season': 2, 'light': 1, 'darkness': 1, 'spring': 1, 'hope': 1, 'winter': 1, 'despair': 1, 'we': 4, 'had': 2, 'everything': 1, 'before': 2, 'us': 2, 'nothing': 1, 'were': 2, 'all': 2, 'going': 2, 'direct': 2, 'to': 1, 'heaven': 1, 'other': 1, 'way': 1, '—': 1, 'in': 2, 'short': 1, 'period': 2, 'so': 1, 'far': 1, 'like': 1, 'present': 1, 'that': 1, 'some': 1, 'its': 2, 'noisiest': 1, 'authorities': 1, 'insisted': 1, 'on': 1, 'being': 1, 'received': 1, 'for': 2, 'good': 1, 'or': 1, 'evil': 1, 'superlative': 1, 'degree': 1, 'comparison': 1, 'only.': 1}
```

<!-- grade_id: words_count_func_1, points: 10, grade: true, solution: false -->
```python
# Test 1: 10 points
assert(word_count_dict['it'] == 10)
```

<!-- grade_id: words_count_func_2, points: 10, grade: true, solution: false -->
```python
# Test 2: 10 points
assert(word_count_dict['the'] == 14)
```

<!-- grade_id: words_count_func_3, points: 10, grade: true, solution: false -->
```python
# Test 3: 10 points (Completely hidden)

```